```python
# Initialize Otter
import otter
grader = otter.Notebook("hw03.ipynb")
```

<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Homework 03: Tables

---

## Textbook References

* [Sequences](https://ccsf-math-108.github.io/textbook/chapters/05/Sequences.html)
* [Tables](https://ccsf-math-108.github.io/textbook/chapters/06/Tables.html)

## Assignment Reminders

- 🚨 Make sure to run the code cell at the top of this notebook that starts with `# Initialize Otter` to load the auto-grader.
- Your Tasks are categorized as auto-graded (📍) and manually graded (📍🔎).
    - For all the auto-graded tasks:
        - Replace the `...` in the provided code cell with your own code.
        - Run the `grader.check` code cell to run some tests on your code.
        - Keep in mind that for homework and project assignments, sometimes there are hidden tests that you will not be able to see the results of that we use for scoring the correctness of your response. **Passing the auto-grader does not guarantee that your answer is correct.**
    - For all the manually graded tasks:
        - You might need to provide your own response to the provided prompt. Do so by replacing the template text "_Type your answer here, replacing this text._" with your own words.
        - You might need to produce a graphic or something else using code. Do so by replacing the `...` in the code cell to generate the image, table, etc.
        - In either case, review the rubric on the associated <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page to understand the scoring criteria.
- Throughout this assignment and all future ones, please be sure to not re-assign variables throughout the notebook! _For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!_
- You may [submit](#Submit-Your-Assignment-to-Canvas) this assignment as many times as you want before the deadline. Your instructor will score the last version you submit once the deadline has passed.
- We encourage you to discuss this assignment with others but make sure to write and submit your own code. Refer to the syllabus to learn more about how to learn cooperatively.

---

## Configure the Notebook

Run the following cell to configure this Notebook.

```python
import numpy as np
from datascience import *
import matplotlib
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

---

## Fruit

To begin this homework, you are going to practice with some of the `Table` methods you learned about this week.

### Task 01 📍

Suppose you have a small fruit stand that has an inventory of 4 apples, 3 oranges, and 1 pineapple.

Create a table called `fruits` that contains this information:
* It should have two columns: `'fruit name'` and `'amount'`.
* It should have three rows, one for each type of fruit.
* The values for `'fruit name'` should be the strings `'apple'`, `'orange'`, and `'pineapple'`
* The values for `'amount'` should be integers `4`, `3`, and `1`.

*Note:* Use lower-case and singular words for the name of each fruit, like `"apple"`.

_Points:_ 2

```python
fruits = ...
fruits
```

```python
grader.check("task_01")
```

---

### Task 02 📍

The file `inventory.csv` contains information about the inventory at a fruit stand.  Each row represents the contents of one box of fruit. Load it as a table named `inventory` using the `Table.read_table()` function. 

**Note**: `Table.read_table(...)` takes one argument (data file name in string format) and returns a table.

_Points:_ 2

```python
inventory = ...
inventory
```

```python
grader.check("task_02")
```

---

### Task 03 📍

Does each box at the fruit stand contain a different fruit? Set `all_different` to `True` if each box contains a different fruit or to `False` if multiple boxes contain the same fruit.

*Hint:* You don't have to write code to calculate the True/False value for `all_different`. Just look at the `inventory` table and assign `all_different` to either `True` or `False` according to what you can see from the table in answering the question.

_Points:_ 2

```python
all_different = ...
all_different
```

```python
grader.check("task_03")
```

---

### Task 04 📍

The file `sales.csv` contains the number of fruit sold from each box last Saturday.  It has an extra column called `price per fruit ($)` that's the price *per item of fruit* for fruit in that box.  The rows are in the same order as the `inventory` table.  Load these data into a table called `sales`.

_Points:_ 2

```python
sales = ...
sales
```

```python
grader.check("task_04")
```

---

### Task 05 📍

How many fruits did the store sell in total on that day (last Saturday)?

_Points:_ 2

```python
total_fruits_sold = ...
total_fruits_sold
```

```python
grader.check("task_05")
```

---

### Task 06 📍

What was the store's total revenue (the total value of all fruits sold) on that day? Assign that value to `total_revenue`.

*Hint:* If you're stuck, think first about how you would compute the total revenue from just the grape sales. It should be `355 * 0.06`, which gives `21.3`.

_Points:_ 2

```python
total_revenue = ...
total_revenue
```

```python
grader.check("task_06")
```

---

### Task 07 📍

Make a new table called `remaining_inventory`.  It should have the same rows and columns as `inventory`, except that the amount of fruit sold from each box should be subtracted from that box's **original** count, so that the `count` column is **updated to be** the amount of fruit remaining after Saturday.

**Notes**:
* As a check, since `355` grapes were sold from the original `517`, then `162` grapes should be in the remaining inventory once you're done.
* We recommend that you do not do the subtractions individually. Let Python (with the help of arrays) do the hard work for you!
* Don't forget that you can always add extra lines of code and intermediate variables. The auto-grader is only checking the value of `remaining_inventory`.

_Points:_ 3

```python
remaining_inventory = ...
remaining_inventory
```

```python
grader.check("task_07")
```

---

## Unemployment


The [Federal Reserve Bank of St. Louis](https://fred.stlouisfed.org/categories/33509) publishes data about jobs in the US. We've imported data from this site on unemployment rates in the United States. There are many ways of defining unemployment rates and our dataset includes two.
1. Non-Employment Index (NEI): This is defined as the percentage of people who are able to work and are looking for a full-time job who can't find a job. 
2. NEI+PTER (Non-Employment Index and 'Part-Time for Economic Reasons'): This is defined as the percentage of people who are able to work and are looking for a full-time job who can't find a job or are only working at a part-time job.
The data that you will use in this assignment contains [quarterly average NEI percentages from 1994 until 2024](https://fred.stlouisfed.org/series/NEIM156SFRBRIC) and [quarterly average NEI+PTER percentages from 1994 until 2024](https://fred.stlouisfed.org/series/NEIPTERM156SFRBRIC)

*In this assignment, you will do some preliminary analysis of these unemployment metrics. In a future assignment, you will work on creating visualizations and summarizing your analysis.*


---

### Task 08 📍

The data are in a CSV file called `unemployment.csv`.  Load the data in that file into a `Table` and assign the table to the name `unemployment`.


_Points:_ 1

```python
unemployment = ...
unemployment
```

```python
grader.check("task_08")
```

---

### Task 09 📍

Dates can be tricky to work with. When you loaded the data into the table `unemployment`, the dates were read as strings. Run the following cell to see this is true.

```python
type(unemployment.column('DATE').item(0))
```

Strings do not always make the best sense to use for dates since we organize dates by months, years, etc. and this is different than the alphabetical order we typically use for strings. A common date type for dates is [`datetime`](https://docs.python.org/3/library/datetime.html). In this class, you are not responsible for knowing how to work directly with the `datetime` format. Instead, we will set up the code to help you with that date type.

**Run the following code cell to convert all the dates in the `'DATE'` column of `unemployment` to `datetime` format.**

_Points:_ 0

```python
# Just run this cell
import datetime
from datetime import datetime, date

def update_date_format(date_string):
    date_object = datetime.strptime(date_string, '%Y-%m-%d').date()
    return date_object

if not isinstance(unemployment.column('DATE').item(0),  date):
    unemployment = unemployment.with_column(
        'DATE', unemployment.apply(update_date_format, 'DATE'))
    print("The date type for the values in the 'DATE' column have been updated.\n")
else:
    print("The date type is correct for the values in the 'DATE' column.\n")
    
display(unemployment)
```

```python
grader.check("task_09")
```

---

### Task 10 📍

For this task:

1. Sort the data in unemployment by NEI in descending order naming the sorted table by_nei.
2. Create another table called by_nei_pter that's sorted by NEI+PTER in descending order instead.

_Points:_ 2

```python
by_nei = ...
by_nei_pter = ...
```

```python
grader.check("task_10")
```

---

### Task 11 📍

For this task:
1. Use `take` to make a table containing the data for the 10 quarters when NEI was greatest.
2. Assign that table to the name `greatest_nei`. Keep in mind that `greatest_nei` should be sorted in descending order of `NEI`.

**Note**: Each row of `unemployment` represents a quarter.


_Points:_ 2

```python
greatest_nei = ...
greatest_nei
```

```python
grader.check("task_11")
```

---

### Task 12 📍

It's believed that many people became PTER (recall: "Part-Time for Economic Reasons") in the "Great Recession" of 2008-2009.  NEI+PTER is the percentage of people who are unemployed (and counted in the NEI) plus the percentage of people who are PTER.  

Using the original `unemployment` table, create an array containing the percentage of people who were PTER in each quarter.  (The first element of the array should correspond to the first row of `unemployment`, and so on.)

_Points:_ 2

```python
pter = ...
pter
```

```python
grader.check("task_12")
```

---

### Task 13 📍

For this task,
1. Add `pter` as a column to `unemployment` (named `"PTER"`).
2. Make sure the data is sorted by the `'DATE'` values in ascending order (`1994-01-01` should be the first date).
3. Assign this table to the name `unemployment` to update the original table.

**Note**: You do not need to do this all in one line of code.

_Points:_ 2

```python
unemployment = ...
unemployment
```

```python
grader.check("task_13")
```

---

### Measuring Relative Change

Unemployment measurements vary, but sometimes the way they change can reflect that something is wrong with the economy. For example, in the following graph, you should see a sharp increase in unemployment rates (NEI+PTER) in 2008.

<img src="./NEI_PTER_over_DATE.png" alt="NEI+PTER values over time" width=400px>

A sudden positive change in unemployment rates as seen around 2008 (the [Great Recession](https://en.wikipedia.org/wiki/Great_Recession))is something to be concerned about.  

There are many benefits of using visualizations like this to detect changes and we will see this in future lessons, but we will concentrate here on detecting dramatic changes numerically.

Run the following code cell to view the first two rows of `unemployment`.

```python
unemployment.show(2)
```

We will look at how the unemployment rates change from quarter to quarter. The first recorded 'NEI+PTER' value on January 1, 1994 was 10.9749 and then on April 1 of that same year the value was 10.6306 reflecting a change of approximately -0.3443. The negative change reflects a decrease in unemployment (an increase in employment). A positive difference would reflect an increase in unemployment (a decrease in employment).

Now, you can't study the amount of change in isolation because a small change during a period of high unemployment could seem dramatic when it really isn't. So, it is better to consider a relative change which is defined as:

$$\text{relative change} = \frac{\text{new value} - \text{old value}}{\text{old value}}$$

This would mean that the relative change in NEI+PTER from January to April, **the first quarter of 1994**, is approximately

$$\frac{10.6306 - 10.9749}{10.9749} \approx	-0.0314.$$

This means there was an approximately 3.14% decrease in the unemployment rate during that quarter (the NEI+PTER went down by 3.14% during that time). Without context, it is difficult to understand if -0.0314 is a common negative value or if it is a significantly large negative value. Later in the course, you will learn how to address this concern more systematically. For now, you will just look for the largest negative relative changes in the data as you complete the next few tasks.

### Task 14 📍

For this task, assign `NEI_PTER_change` to an array with all of the relative changes in NEI+PTER values as we just described.
* The `NEI_PTER_change` array should contain 119 items.
* The first value in `NEI_PTER_change` should be approximately `-0.0314`.
  
**Notes**: 
* We've provided some optional intermediate variables to help you complete this task. You can delete them if you don't end up using them. Just make sure you have `NEI_PTER_change` defined as requested in your code.
* Notice that when you are calculating the relative change values, you'll never divide by the last value in the data set. For now, we recommend using the table method [`exclude`](https://datascience.readthedocs.io/en/master/_autosummary/datascience.tables.Table.exclude.html#datascience.tables.Table.exclude) to handle this. As a hint, `unemployment.exclude(0)` will create a table without the first row of `unemployment`.
* You might find the `np.diff` function helpful for this task.

_Points:_ 2

```python
NEI_PTER_array = ...
first_119_NEI_PTER_array = ...
NEI_PTER_change = ...
NEI_PTER_change
```

```python
grader.check("task_14")
```

Run the following code to add the values in the `NEI+PTER_change` array to the table under the column label of `'NEI+PTER CHANGE'`. (You'll work with `np.append` later in this course.)

```python
# Run this code after completing Task 14
unemployment_with_change = unemployment.with_column('NEI+PTER CHANGE',
                                                    np.append(make_array(0), NEI_PTER_change))
unemployment_with_change
```

---

### Defining Significance

When dealing with large-scale economic indicators like unemployment, even small percentage changes can represent substantial shifts in the labor market. There are many ways to define when numerical values are significant. Sometimes significance is determined by some, seemingly arbitrary, threshold value.

For the sake of this assignment, let's say that a 10% increase in NEI+PTER measurements reflects a significant (bad) change in unemployment. In other words, **consider any relative change more than 0.1 to be significant.**

### Task 15 📍

For this task, create a table called `significant_changes` that contains only the rows from the `unemployment_with_change` table with 'NEI+PTER CHANGE' more than 0.1.

_Points:_ 3

```python
...
```

```python
grader.check("task_15")
```

---

### Task 16 📍🔎

<!-- BEGIN QUESTION -->

What do you notice about the timing of these significant negative changes in NEI+PTER? In your response:
* Highlight any global events that coincide with the dates found.
* Provide your opinion about how the events you highlighted might impact the NEI+PTER changes.

_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

## Submit Your Assignment to Canvas

Follow these steps to submit your homework assignment:

1. **Review the Rubric:** View the rubric on the associated Canvas Assignment page to understand the scoring criteria.
2. **Run the Auto-Grader:** Ensure you have executed the code cell containing the command `grader.check_all()` to run all tests for auto-graded tasks marked with 📍. This command will execute all auto-grader tests sequentially.
3. **Complete Manually Graded Tasks:** Verify that you have responded to all the manually graded tasks marked with 📍🔎.
4. **Save Your Work:** In the notebook's Toolbar, go to `File -> Save Notebook` to save your work and create a checkpoint.
5. **Download the Notebook:** In the notebook's Toolbar, go to `File -> Download IPYNB` to download the notebook (`.ipynb`) file.
6. **Upload to Canvas:** On the Canvas Assignment page, click "Start Assignment" or "New Attempt" to upload the downloaded `.ipynb` file.

---

## Attribution

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a> and derived from the <a href="https://www.data8.org/">Data 8: The Foundations of Data Science</a> offered by the University of California, Berkeley.

<img src="./by-nc-sa.png" width=100px>

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for task_01 (points: 2)
```
>>> isinstance(fruits, Table)
True
```
failure_message: ❌ fruits should be a Table. Try using Table().with_columns(...) to build the table.
success_message: ✅ fruits is a Table.
```
>>> fruits.sort(0)
fruit name | amount
apple      | 4
orange     | 3
pineapple  | 1
```
failure_message: ❌ fruits doesn't look correct. Double-check the capitalization, spelling, and counts.
success_message: ✅ fruits looks correct!


Tests for task_02 (points: 2)
```
>>> isinstance(inventory, Table)
True
```
failure_message: ❌ inventory should be a Table.
success_message: ✅ inventory is a Table.
```
>>> inventory.sort(0)
box ID | fruit name | count
25274  | apple      | 20
26187  | strawberry | 255
43566  | peach      | 40
48800  | orange     | 35
52357  | strawberry | 102
53686  | kiwi       | 45
57181  | strawberry | 123
57930  | grape      | 517
```
failure_message: ❌ inventory doesn't look correct.
success_message: ✅ inventory looks correct.


Tests for task_03 (points: 2)
```
>>> isinstance(all_different, bool)
True
```
failure_message: ❌ all_different should be assigned to True or False. Double check your spelling/capitalization of True or False.
success_message: ✅ all_different is assigned to True or False.


Tests for task_04 (points: 2)
```
>>> isinstance(sales, Table)
True
```
failure_message: ❌ sales should be a Table.
success_message: ✅ sales is a Table.
```
>>> sales.sort(0)
box ID | fruit name | count sold | price per fruit ($)
25274  | apple      | 0          | 0.8
26187  | strawberry | 25         | 0.15
43566  | peach      | 17         | 0.8
48800  | orange     | 35         | 0.6
52357  | strawberry | 102        | 0.25
53686  | kiwi       | 3          | 0.5
57181  | strawberry | 101        | 0.2
57930  | grape      | 355        | 0.06
```
failure_message: ❌ sales doesn't look correct.
success_message: ✅ sales looks correct.


Tests for task_05 (points: 2)
```
>>> isinstance(total_fruits_sold, (int, np.integer))
True
```
failure_message: ❌ total_fruits_sold should be an integer.
success_message: ✅ total_fruits_sold is an integer.
```
>>> total_fruits_sold > 10
True
```
failure_message: ❌ total_fruits_sold is not a possible quantity.
success_message: ✅ total_fruits_sold is a possible quantity.


Tests for task_06 (points: 2)
```
>>> 50 <= total_revenue <= 150
True
```
failure_message: ❌ total_revenue is not a possible value. It should be a number between 50 and 150.
success_message: ✅ total_revenue is a possible value.
```
>>> total_revenue == 21.3
False
```
failure_message: ❌ total_revenue seems to reflect the revenue for grapes. You also need to include the revenue for all the other fruits sold.
success_message: ✅ total_revenue doesn't just reflect the revenue for grapes.


Tests for task_07 (points: 3)
```
>>> isinstance(remaining_inventory, Table)
True
```
failure_message: ❌ remaining_inventory should be a Table.
success_message: ✅ remaining_inventory is a Table.
```
>>> remaining_inventory.where(1, 'grape')
box ID | fruit name | count
57930  | grape      | 162
```
failure_message: ❌ The updated grape inventory doesn't look correct.
success_message: ✅ The updated grape inventory looks good.


Tests for task_08 (points: 1)
```
>>> isinstance(unemployment, Table)
True
```
failure_message: ❌ It doesn't seem that you have assigned the name unemployment to a table. Try using the command Table().read_table(...)
success_message: ✅ It seems that you have assigned the name unemployment to a table.
```
>>> np.isclose(np.sum(unemployment.column('NEI')), 1044.46519909577)
True
```
failure_message: ❌ The contents in unemployment don't seem correct. Double-check that you are using the unemployment.csv
success_message: ✅ The contents in unemployment seem correct.


Tests for task_09 (points: 0)
```
>>> isinstance(unemployment.column('DATE').item(0), date)
True
```
failure_message: ❌ It doesn't seem like you've run the cell to convert the date data types. Make sure to run the above code cell.
success_message: ✅ You've converted all the dates.


Tests for task_10 (points: 2)
```
>>> by_nei.take(0).column(0).item(0) == date(2020, 4, 1)
True
```
failure_message: ❌ It doesn't look like you have the correct values in by_nei.
success_message: ✅ It seems like you have the correct values in by_nei.
```
>>> by_nei_pter.take(0).column(0).item(0) == date(2020, 4, 1)
True
```
failure_message: ❌ It doesn't look like you have the correct values in by_nei_pter.
success_message: ✅ It seems like you have the correct values in by_nei_pter.


Tests for task_11 (points: 2)
```
>>> greatest_nei.num_rows == 10
True
```
failure_message: ❌ There are not 10 rows (quarters) in your table.
success_message: ✅ There are 10 rows (quarters) in your table.


Tests for task_12 (points: 2)
```
>>> isinstance(pter, np.ndarray)
True
```
failure_message: ❌ pter doesn't seem to be an array.
success_message: ✅ pter is an array.
```
>>> round(pter.item(6), 4) != -1.1446
True
```
failure_message: ❌ Be careful of subtracting the arrays in the wrong order.
success_message: ✅ You didn't seem to subtract the correct arrays in the wrong order.


Tests for task_13 (points: 2)
```
>>> unemployment.num_columns == 4
True
```
failure_message: ❌ Your table doesn't have the correct number of columns.
success_message: ✅ Your table has the correct number of columns.
```
>>> np.all(unemployment.sort('DATE').column('DATE') == unemployment.column('DATE'))
True
```
failure_message: ❌ Your table doesn't appear to be sorted, in ascending order, by the date values.
success_message: ✅ Your table appears to be sorted, in ascending order, by the date values.


Tests for task_14 (points: 2)
```
>>> isinstance(NEI_PTER_change, np.ndarray)
True
```
failure_message: ❌ NEI_PTER_change should represent an array.
success_message: ✅ NEI_PTER_change represents an array.
```
>>> len(NEI_PTER_change)
119
```
failure_message: ❌ NEI_PTER_change doesn't have the correct number of items.
success_message: ✅ NEI_PTER_change has the correct number of items.


Tests for task_15 (points: 3)
```
>>> isinstance(significant_changes, Table)
True
```
failure_message: ❌ significant_changes should be a Table.
success_message: ✅ significant_changes is a Table.
