```python
# Initialize Otter
import otter
grader = otter.Notebook("lab07.ipynb")
```

<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Lab 07 - Testing Hypotheses

## References

* [Sections 11.0 - 11.4 of the Textbook](https://ccsf-math-108.github.io/textbook/chapters/11/Testing_Hypotheses.html)
* [datascience Documentation](https://datascience.readthedocs.io/)
* [Code Reference](https://ccsf-math-108.github.io/materials-sp25/resources/code-reference.html)

---

## Lab Assignment Reminders

- 🚨 Make sure to run the code cell at the top of this notebook that starts with `# Initialize Otter` to load the auto-grader.
- Your tasks are categorized as auto-graded (📍) and manually graded (📍🔎):
    - **For all auto-graded tasks:**
        - Replace the `...` in the provided code cell with your own code.
        - Run the `grader.check` code cell to execute tests on your code.
        - There are no hidden auto-grader tests in the lab assignments. This means if you pass the tests, you can assume you've completed the task successfully.
    - **For all manually graded tasks:**
        - You may need to provide your own response to the provided prompt. Replace the template text "_Type your answer here, replacing this text._" with your own words.
        - You might need to produce a graphic or another output using code. Replace the `...` in the code cell to generate the image, table, etc.
        - In either case, check your response with a classmate, a tutor, or the instructor before moving on.
- Throughout this assignment and all future ones, please **do not re-assign variables** throughout the notebook! _For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you may fail tests that you thought you were passing previously!_
- You may [submit](#Submit-Your-Assignment-to-Canvas) this assignment as many times as you want before the deadline. Your instructor will score the last version you submit once the deadline has passed.
- **Collaborating on labs is encouraged!** You should rarely remain stuck for more than a few minutes on questions in labs, so ask an instructor or classmate for help. (Explaining things is beneficial, too -- the best way to solidify your knowledge of a subject is to explain it.) However, please don't just share answers.

---

## Configure the Notebook

Run the following cell to configure this Notebook.

```python
from datascience import *
import numpy as np
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

---

## Therapeutic Touch

[Therapeutic touch (TT)](https://en.wikipedia.org/wiki/Therapeutic_touch) is an alternative medicine practice that is based on the idea that the human body has an energy field (called the human energy field, or HEF) that can be sensed, manipulated and balanced to promote health and relaxation. Practitioners of TT lightly place their hands on or above a patient's body to sense the patient's energy and then send healthy energy back.

---

## Emily Rosa

In 1996, [Emily Rosa](https://en.wikipedia.org/wiki/Emily_Rosa) was a 4th grade student who was very familiar with the world of TT, thanks to her parents, who were both medical practitioners and skeptics of TT. For her 4th grade science fair project, she designed an experiment that tested whether or not TT practitioners could truly sense a person's HEF. Her work was published in the Journal of the American Medical Association in 1998 making her the youngest person to have research published in a peer reviewed medical journal.

---

## Emily's Experiment

Emily's experiment was clean, simple, and effective. Due to her parents' occupations in the medical field, she had wide access to people who were TT practitioners.

Emily's experiment involved 21 TT practitioners and they were each given 10 trials. Each trial involved the practitioner extending both of their hands through a screen which they couldn't see through. On the other side of the screen, Emily placed her own hand near one of the practitioner's hands - either right or left based on the flip of a fair coin. She then asked the practitioner to tell her which of their hands detected her HEF. The idea is that if a practitioner can truly interact with a person's HEF, they should be able to answer this question correctly.

Overall, through the 210 trials, the practitioner picked the correct hand 44% of the time.

Emily's main goal here was to test whether or not the TT practitioners' guesses were random, like the flip of a coin. In many medical experiments, this is the norm. We test whether or not the treatment has an effect, not necessarily whether or not the treatment actually works.

We will now begin to formulate this experiment in terms of the terminology we learned in this course.

---

## Task 01 📍🔎

<!-- BEGIN QUESTION -->

Emily suspected that the practitioners could not sense a person's HEF and were merely guessing at the correct hand. Using this idea, describe Emily's model for how likely the TT practitioners are to choose the correct hand.

What is the alternative model to hers (i.e. the model hers is meant to discredit)?

Discuss with students around you to come to a conclusion. Check-in with a tutor or the instructor if you are stuck. Make sure you have this right before moving on.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

## Task 02 📍

Remember that the practitioner got the correct answer 44% (0.44) of the time. According to Emily's model, on average, what proportion of times do we expect the practitioner to guess the correct hand? Make sure your answer is between 0 and 1. 


```python
expected_proportion_correct = ...
expected_proportion_correct
```

```python
grader.check("task_02")
```

The goal now is to see if the deviation in the experimental results from this expected proportion is due to something other than chance.

---

## Task 03 📍

We usually use a statistic to help determine which model the evidence points towards. What is a statistic that we can use to compare outcomes under Emily’s model to what was observed? Assign `valid_stat` to an array of integer(s) representing test statistics that Emily can use: 

1. The difference between the expected percent correct and the actual percent correct
2. The absolute difference between the expected percent correct and the actual percent correct
3. The sum of the expected percent correct and the actual percent correct



```python
valid_stat = ...
valid_stat
```

```python
grader.check("task_03")
```

---

## Task 04 📍🔎

<!-- BEGIN QUESTION -->

Why is the statistic from Task 03 the best choice for comparing outcomes in Emily's experiment? How does it relate to the models you defined in Task 01?

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

## Task 05 📍

Define the function `statistic` which takes in an expected proportion and an actual proportion, and returns the value of the statistic chosen in Task 3. Assume that the argument takes in proportions, but  return your answer as a percentage. 

*Hint:* Remember we are asking for a **percentage**, not a proportion. 


```python
def statistic(expected_prop, actual_prop):
    ...
```

```python
grader.check("task_05")
```

---

## Task 06 📍

Use your newly defined function to calculate the observed statistic from Emily's experiment. 


```python
observed_statistic = ...
observed_statistic
```

```python
grader.check("task_06")
```

> **Is this observed statistic consistent with what we might see under Emily’s model?**

In order to answer this question, we must simulate the experiment as though Emily's model was correct, and calculate our statistic for every simulation.

---

## `sample_proportions`

`sample_proportions` can be used to randomly sample from multiple categories when you know the proportion of data points that are expected to fall in each category. `sample_proportions` takes two arguments: the sample size and an array that contains the distribution of categories in the population (should sum to 1).

Consider flipping a fair coin, where the two outcomes (coin lands heads and coin lands tails) occur with an equal chance. We expect that half of all coin flips will land heads, and half of all coin flips will land tails.

Run the following cell to see the simulation of 10 flips of a fair coin. Let the first item of `coin_proportions` be the proportion of heads and the second item of `coin_proportions` be the proportion of tails.

*Observe what happens when you run this cell multiple times. The proportion of coin flips that land heads and tails appears to change as you are simulating flipping 10 coins each time!*

```python
coin_proportions = make_array(0.5, 0.5) 
ten_flips = sample_proportions(10, coin_proportions)
ten_flips
```

`sample_proportions` returns an array that is the same length as the proportion array that is passed through. It contains the proportion of each category that appears in the sample. 

In our example, the first item of `ten_flips` is the simulated proportion of heads and the second item of `ten_flips` is the simulated proportion of tails.

```python
simluated_proportion_heads = ten_flips.item(0)
simluated_proportion_tails = ten_flips.item(1)

print("In our simluation, " + str(simluated_proportion_heads) + " of flips were heads and " \
      + str(simluated_proportion_tails) + " of flips were tails.")
```

---

## Task 07 📍

To begin simulating, we should start by creating a representation of Emily's model to use for our simulation. This will be an array with two items in it. The first item should be the proportion of times, assuming that Emily’s model was correct, a TT practictioner picks the correct hand. The second item should be the proportion of times, under the same assumption, that the TT practitioner picks the incorrect hand. Assign `model_proportions` to this array. 

After this, we can simulate 210 hand choices, as Emily evaluated in real life, and find a single statistic to summarize this instance of the simulation. Use the `sample_proportions` function and assign the proportion of correct hand choices (out of 210) to `simulation_proportion_correct`. Lastly, use your statistic function to assign `one_statistic`  to the value of the statistic for this one simulation.

*Hint:* `sample_proportions` usage can be found [here](http://data8.org/su19/python-reference.html).


```python
# This saves the random state of our code so that we can 
# generate the same numbers each time we run the code.
# Please do not change this next line. 
np.random.seed(1234)

# Your work goes below this comment.
model_proportions = ...
simulation_proportion_correct = ...
one_statistic = ...
one_statistic
```

```python
grader.check("task_07")
```

---

## Task 08 📍

Let's now see what the distribution of statistics is actually like under Emily's model. 

1. Define the function `simulation_and_statistic` to take in the `model_proportions` array and the expected proportion of times a TT practitioner would guess a hand correctly under Emily's model. The function should simulate Emily running through the experiment 210 times and return the statistic of this one simulation. _This should follow the same pattern as the code you did in the previous problem._
2. Using this function, assign `simulated_statistics` to an array of 1000 statistics that you calculated under the assumption that Emily's model was true.

```python
def simulation_and_statistic(model_proportions, expected_proportion_correct):
    '''Simulates 210 TT hand choices under Emily’s model. 
    Returns one statistic from the simulation.'''
    ...
```

```python
num_repetitions = 1000

simulated_statistics = ...

for ... in ...:
    ...

```

```python
grader.check("task_08")
```

---

Let's view the distribution of the simulated statistics under Emily's model, and visually compare where the observed statistic lies relative to the simulated statistics.

```python
t = Table().with_column('Simulated Statistics', simulated_statistics)
t.hist()
plt.scatter(observed_statistic, 0, color='red', s=30, zorder=3);
```

We can make a visual argument as to whether we believe the observed statistic is consistent with Emily’s model. Here, since larger values of the test statistic suggest the alternative model (where the chance of guessing the correct hand is something other than 50%), we can formalize our analysis by finding what proportion of simulated statistics were as large or larger than our observed test statistic (the area at or to the right of the observed test statistic). If this area is small enough, we’ll declare that the observed data are inconsistent with our simulated model.

---

## Task 09 📍

Calculate the proportion of simulated statistics greater than or equal to the observed statistic. 

You might want to use `np.count_nonzero` for this.


```python
proportion_greater_or_equal = ...
proportion_greater_or_equal
```

```python
grader.check("task_09")
```

By convention, we often compare the proportion we just calculated to 0.05. If the proportion of simulated statistics greater than or equal to the observed statistic is sufficiently small (less than or equal to 0.05), then this is evidence against Emily's model. Otherwise, we don’t have any reason to doubt Emily’s model. 

This should help you make your own conclusions about Emily Rosa's experiment. 

Therapeutic touch fell out of use after this experiment, which was eventually accepted into one of the premier medical journals. TT practitioners hit back and accused Emily and her family of tampering with the results, while some claimed that Emily's bad spiritual mood towards therapeutic touch made it difficult to read her HEF. Whatever it may be, Emily's experiment is a classic example about how anyone, with the right resources, can test anything they want!


---

## Task 10 📍🔎

<!-- BEGIN QUESTION -->

Now, take some time to reflect on the questions below and then, discuss with your peers.

1. Is the data more consistent with Emily's model (practioners were randomly guessing)?
2. What does this mean in terms of Emily's experiment? Do the TT practitioners' answers follow an even chance model or is there something else at play? 

Write you response to these questions in the following cell.

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

## Submit Your Assignment to Canvas

Follow these steps to submit your lab assignment:

1. **Check the Assignment Completion Requirements:** This assignment is scored as Complete or Incomplete. Make sure to check with your instructor about their requirements for a Complete score. 
2. **Run the Auto-Grader:** Ensure you have executed the code cell containing the command `grader.check_all()` to run all tests for auto-graded tasks marked with 📍. This command will execute all auto-grader tests sequentially.
3. **Complete Manually Graded Tasks:** Verify that you have responded to all the manually graded tasks marked with 📍🔎.
4. **Save Your Work:** In the notebook's Toolbar, go to `File -> Save Notebook` to save your work and create a checkpoint.
5. **Download the Notebook:** In the notebook's Toolbar, go to `File -> Download HTML` to download the HTML version (`.html`) of this notebook.
6. **Upload to Canvas:** On the Canvas Assignment page, click "Start Assignment" or "New Attempt" to upload the downloaded `.html` file.

---

## Attribution

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a> and derived from the <a href="https://www.data8.org/">Data 8: The Foundations of Data Science</a> offered by the University of California, Berkeley.

<img src="./by-nc-sa.png" width=100px>

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for task_02 (points: None)
```
>>> 0 <= expected_proportion_correct <= 1
True
```
failure_message: ❌ Remember that a propbability value is a number between 0 and 1, inclusive.
success_message: ✅ Great work. A propbability value is a number between 0 and 1, inclusive.
```
>>> expected_proportion_correct == 0.5
True
```
failure_message: ❌ Emily believes they have a 50% chance of guessing correctly.
success_message: ✅ You got it!


Tests for task_03 (points: None)
```
>>> type(valid_stat) == np.ndarray
True
```
failure_message: ❌ valid_stat should be an array. Try using make_array().
success_message: ✅ valid_stat is an array.
```
>>> int(sum(valid_stat)) == 2
True
```
failure_message: ❌ Remember that she wants to measure whether or not the treatment has an effect, not whether or not the treatment actually works.
success_message: ✅ Great work. Your answers seem correct!


Tests for task_05 (points: None)
```
>>> int(round(statistic(0.5, 0.5) + statistic(0.4, 0.1), 1)) == 30
True
```
failure_message: ❌ We tested your function for a few values it doesn't seem to be working. Try subracting the inputs, computing the absolute value, and converting the proportion into a percent value.
success_message: ✅ We tested your function for a few values it seems to be working.
```
>>> int(statistic(0.4, 0.1) - statistic(0.1, 0.4)) == 0 and abs(statistic(0.4, 0.1)) == abs(statistic(0.1, 0.4))
True
```
failure_message: ❌ Don't forget to use the absolute value.
success_message: ✅ Great work using absolute value.


Tests for task_06 (points: None)
```
>>> int(round(observed_statistic, 2)) == 6
True
```
failure_message: ❌ Your value doesn't seem correct. Remember that the observed proportion was 0.44.
success_message: ✅ Your value seems correct!


Tests for task_07 (points: None)
```
>>> len(model_proportions) == 2
True
```
failure_message: ❌ model_proportions should have 2 items.
success_message: ✅ model_proportions has 2 items.
```
>>> len(np.unique(model_proportions)) == 1
True
```
failure_message: ❌ Your values should be the same. Revisit Emily's model.
success_message: ✅ Your values are the same.
```
>>> sum(model_proportions) == 1
True
```
failure_message: ❌ Your values should sum to be 1.
success_message: ✅ Your values sum to be 1.
```
>>> type(simulation_proportion_correct) == float
True
```
failure_message: ❌ simulation_proportion_correct should be a float. Don't for get to get the first item from the array.
success_message: ✅ simulation_proportion_correct is a float.
```
>>> round(simulation_proportion_correct, 2)
0.48
```
failure_message: ❌ This value for simulation_proportion_correct doesn't seem correct. Double check your sample size and the population proportions.
success_message: ✅ This value for simulation_proportion_correct seems correct.
```
>>> np.isclose(one_statistic, 1.9, atol=0.01)
True
```
failure_message: ❌ This value for one_statistic doesn't seem correct. Try using the statistic function.
success_message: ✅ This value for one_statistic seems correct.


Tests for task_08 (points: None)
```
>>> callable(simulation_and_statistic)
True
```
failure_message: ❌ simulation_and_statistic should be a function.
success_message: ✅ You defined a function called simulation_and_statistic.
```
>>> len(simulated_statistics) == 1000
True
```
failure_message: ❌ There should be 1000 values in simulated_statistics.
success_message: ✅ It seems that you've stored 1000 values in simulated_statistics.
```
>>> len(set(simulated_statistics)) > 1
True
```
failure_message: ❌ simulated_statistics does not contain more than 1 value.
success_message: ✅ simulated_statistics contains more than 1 value.
```
>>> np.all(simulated_statistics <= 30)
True
```
failure_message: ❌ All of the values in your simulated_statistics array should be at most 30.
success_message: ✅ All of the values in your simulated_statistics array are at most 30.
```
>>> np.all(simulated_statistics >= 0)
True
```
failure_message: ❌ All of the values in your simulated_statistics array should be at least 0.
success_message: ✅ All of the values in your simulated_statistics array are at least 0.
```
>>> 0 <= simulation_and_statistic(model_proportions, expected_proportion_correct) <= 25
True
```
failure_message: ❌ Your value for simulation_and_statistic should be between 0 and 25, inclusive.
success_message: ✅ Great. Your value for simulation_and_statistic is between 0 and 25, inclusive. This is what we expected.


Tests for task_09 (points: None)
```
>>> 0 <= proportion_greater_or_equal <= 1
True
```
failure_message: ❌ proportion_greater_or_equal should be a number between 0 and 1, inclusive.
success_message: ✅ proportion_greater_or_equal is a possible value.
```
>>> proportion_greater_or_equal * 1000 == np.count_nonzero(simulated_statistics >= observed_statistic)
True
```
failure_message: ❌ Your value for proportion_greater_or_equal doesn't seem correct. Check when the simulated statistic is greater than or equal the observed statistic. Divide the number of times when that happened by the number of simulations you performed.
success_message: ✅ Your value for proportion_greater_or_equal seems correct.
