```python
# Initialize Otter
import otter
grader = otter.Notebook("hw05.ipynb")
```

<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Homework 5: Data Analysis

## Textbook References

* [Chapter 8 - Functions and Tables](https://ccsf-math-108.github.io/textbook/chapters/08/Functions_and_Tables.html#) 
* [`datascience` Documentation](https://datascience.readthedocs.io/)

---

## Assignment Reminders

- 🚨 Make sure to run the code cell at the top of this notebook that starts with `# Initialize Otter` to load the auto-grader.
- Your Tasks are categorized as auto-graded (📍) and manually graded (📍🔎).
    - For all the auto-graded tasks:
        - Replace the `...` in the provided code cell with your own code.
        - Run the `grader.check` code cell to run some tests on your code.
        - Keep in mind that for homework and project assignments, sometimes there are hidden tests that you will not be able to see the results of that we use for scoring the correctness of your response. **Passing the auto-grader does not guarantee that your answer is correct.**
    - For all the manually graded tasks:
        - You might need to provide your own response to the provided prompt. Do so by replacing the template text "_Type your answer here, replacing this text._" with your own words.
        - You might need to produce a graphic or something else using code. Do so by replacing the `...` in the code cell to generate the image, table, etc.
        - In either case, review the rubric on the associated <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page to understand the scoring criteria.
- Throughout this assignment and all future ones, please be sure to not re-assign variables throughout the notebook! _For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!_
- You may [submit](#Submit-Your-Assignment-to-Canvas) this assignment as many times as you want before the deadline. Your instructor will score the last version you submit once the deadline has passed.
- We encourage you to discuss this assignment with others but make sure to write and submit your own code. Refer to the syllabus to learn more about how to learn cooperatively.

---

## Configure the Notebook

Run the following cell to configure this Notebook.

```python
from datascience import *
import numpy as np
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
```

---

## 🌉 San Francisco City and County Employee Salaries

In this homework assignment, you will combine several of the skills and concepts you've learned so far to analyze [Employee Compensation](https://data.sfgov.org/City-Management-and-Ethics/Employee-Compensation/88g8-5mnd) data provided by the SF Controller's Office. We have filtered it to retain just the relevant columns and restricted the data to the fiscal year 2022-2023 (July 1, 2022 - June 30, 2023). 

Run the following cell to load that data into a table called `sf`.

```python
sf = Table.read_table('sf_fy22_23.csv')
sf
```

---

### 📝 Data Notes

* Each line represents an employee's job information such as job family, salary, and benefits.
* The San Francisco Controller's Office  provides [a PDF explaining what each variable means](https://data.sfgov.org/api/views/88g8-5mnd/files/OMBVvreoXRjXG6oP4Ts4497dNxt14XlBqB6uIL6cq-o?download=True&filename=N:\EIS\DataCoordination\Metadata%20Spring%20Cleaning\CON_DataDictionary_Employee-Compensation.pdf).
There are a few interesting details in the data set based on [Employee Compensation - FAQ](https://sfdigitalservices.gitbook.io/dataset-explainers/employee-compensation-faq):
    * Employees have been given an employee identifier for this data and is not connected to their actual employee identification.
    * An employee can appear on multiple lines.
    * Rows may have zero dollar amounts.
    * Adjustments to employee's pay can result in negative dollar amounts.

---

### 🔔 Table Size Warning

This table has over 40 thousand rows in it. Just loading the data takes up more memory than most of our data sets so far. This means that you want to be careful with creating too many copies of the table since you have a limited amount of memory. We've set up the assignment, so you can complete the tasks with the provided amount of memory, but you may accidentally crash the kernel by running out of memory.

**If your kernel keeps crashing because of running out of memory:**
* Make sure you are only running this notebook and all your other kernels are shut down.
* Talk with a staff member about things you might be doing that utilize a lot of resources.

---

### Task 01 📍

One employee might appear on multiple lines per year if they changed positions and/or departments mid-year, which happens with some frequency. This may or may not be true for this particular fiscal year. There are 41,295 rows in the `sf` table. Determine how many unique employee identifiers there are in the `sf` table to see if there are any employees listed more than once. Assign `num_employee_ids` to the number (`int`) of unique employee identifiers in the `sf` table.

*Hint: Consider using the `group` table method.*

_Points:_ 2

```python
num_employee_ids = ...
num_employee_ids
```

```python
grader.check("task_01")
```

---

### Task 02 📍🔎

<!-- BEGIN QUESTION -->

There are a few departments that have a large number of employees. For this task, produce a bar chart that shows the five largest departments in the `sf` table ranked by employee count.

_Points:_ 2

```python
...

# Leave this to provide a title for your visualization
plt.title('Top 5 Departments By Employee Count')
plt.show()
```

<!-- END QUESTION -->

---

### Task 03 📍

Some employees for the city and county make a lot of money. Identify the job title that makes the most in terms of total compensation (the sum of salary and benefits).
1. Assign the largest total compensation (a `float` value) to `max_compensation`.
2. Assign the job title (a `str` value) associated with the highest total compensation to the name `max_compensation_position`.

_Points:_ 2

```python
max_compensation = ...
max_compensation_position = ...
print(f"The largest total compensation is ${max_compensation:,}")
print(f"The position associated with this compensation is: {max_compensation_position}")
```

```python
grader.check("task_03")
```

---

### Task 04 📍🔎

<!-- BEGIN QUESTION -->

Next, summarize the distribution of total compensation values so you can compare that large salary with the salaries of the rest of the employees. Specifically, create a histogram showing the distribution of total compensation.
* Use the `unit = '$'` argument
* Use the default bins.

_Points:_ 2

```python
...

# Leave this to provide a title for your visualization
plt.title('Total Compensation')
plt.show()
```

<!-- END QUESTION -->

---

### Task 05 📍🔎

<!-- BEGIN QUESTION -->

In the provided Markdown cell, explain why there seems to be no information visualized on the right side of the histogram created in Task 04. Make sure you connect your response to the `max_compensation` value you calculated.

_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

---

### Task 06 📍

Your goal for this task is to create a table titled `top_5_departments` featuring two columns: `department` listing department names and `total_compensation median` displaying the median total compensation for each department. Once done, the final table should showcase the 5 departments in the data set with the highest median total compensation values.

_Points:_ 3

```python
top_5_departments = ...
top_5_departments
```

```python
grader.check("task_06")
```

---

### Task 07 📍

Define a function called `total_compensation_range` that takes an array of compensation values as input and returns the range, calculated as the difference between the highest and lowest total compensation within the department for that organization group:

$$\text{Total Compensation Range} = \max(\text{Total Compensations}) - \min(\text{Total Compensations})$$

_Points:_ 2

```python
...
```

```python
grader.check("task_07")
```

---

### Task 08 📍

Create a table called `department_ranges` where each row represents a department and each column represents an organization group. The table's values should indicate the total compensation range for each combination of department and organization group.

**Hint**: What table function allows you to specify the rows and columns of a new table?

_Points:_ 3

```python
department_ranges = ...
department_ranges
```

```python
grader.check("task_08")
```

---

### Overtime Ratio

California has laws in place to help govern how much an employee should be paid for overtime work. Visit the [Labor Commissioner's Office FAQ about Overtime](https://www.dir.ca.gov/dlse/faq_overtime.htm) if you are curious about these laws. There could be several reasons for why an employee works overtime. Some are healthy reasons and some are not. The ratio of overtime compensation to total compensation can provide a signal for health of departments and their employees. According to [Indeed.com](https://www.indeed.com/career-advice/career-development/working-overtime), here are some potential disadvantages of working extra hours:

> **Focus loss**: You will likely want to take breaks while working overtime, and you may lose focus and productivity naturally as your working hours increase. 
> 
> **Safety and health risk**: Working longer hours also can be dangerous, depending on the job. Working overtime regularly can also disrupt your work-life balance, lead to burnout or create health risks, such as sitting at a computer for long periods. Due to these risks, more companies are limiting the number of hours worked in certain positions, such as truck drivers. 
> 
> **Less work-life balance**: There are only 24 hours in the day, and working overtime reduces the time for a good work-life balance. More work hours mean fewer hours for family, relaxation and sleep.

On the other hand, here are some potential advantages for working extra hours:

> **Extra money**: Working overtime means you earn extra money for extra hours, usually 1.5 to two times your hourly rate. Overtime pay could allow you to bulk up your salary and have more money for living expenses, vacations or savings. 
>
>**Career boost**: Working overtime may also help advance your career—and potential earnings—by proving you’re reliable and motivated in your role. Volunteering to work extra hours shows your current employer that you’re dedicated to doing a good job. 
>
>**Gain experience**: Overtime is an opportunity to gain job experience and polish your skills. Typically, the more you work at something, the better you become.

One way to analyze overtime work for San Francisco employees is to consider the overtime pay values in the data set, but these values are relative to the amount of money the employee makes. To account for this you could consider overtime pay for an employee as a quantity relative to that employee's salary. This would allow you to compare these ratios for all the employees. Specifically, use the measurement called the overtime ratio defined as:

$$\text{Overtime Ratio} = \frac{\text{Overtime Pay}}{\text{Salary}}$$

There are many special cases, when it comes to overtime pay, that might make analyzing this ratio in general challenging. So, to keep things more simple, narrow your focus to employees that make at least a certain salary. Let that base salary reflect working 40 hrs per week for 52 weeks at the minimum wage of the time, 16.99 dollars per hour.

---

### Task 09 📍

For this task, create a table called `sf_above_35k` that contains the same information in the `sf` table and overtime ratios for employees making above $35,000. Make sure that `sf_above_35k` has:
 * Only employees from the `sf` table that have salaries above 35,000 dollars AND
 * The same columns as the `sf` table with an additional column `overtime_ratio` containing the overtime ratios from the above formula.

_Points:_ 3

```python
overtime_ratios = ...

sf_above_35k = ...
sf_above_35k
```

```python
grader.check("task_09")
```

---

### Task 10 📍🔎

<!-- BEGIN QUESTION -->

Is overtime collected similarly over the various salaries in the `sf_above_35k` table? See if there is any relationship between salary and the overtime ratios. For this task, create a visualization showing the relationship between salary and overtime ratio values for employees who make above $35,000.

_Points:_ 2

```python
...

# Leave this to provide a title for your visualization
plt.title('Overtime Ratio vs. Salary')
plt.show()
```

<!-- END QUESTION -->

---

### Task 11 📍🔎

<!-- BEGIN QUESTION -->

Take a close look at the scatter plot of Overtime Ratio vs. Salary. You may notice a few interesting patterns:
* A noticeable concentration of high overtime ratios for salaries around \$150,000.
* A sharp decline in overtime ratios for those with salaries above \$200,000.

In the following markdown cell, reflect on this graphic by addressing the following:
1. Why do you think there is a bump in overtime ratio around \$150,000? What types of jobs or organizational policies might explain this?
2. Why do higher-paid employees (above \$200,000) have lower overtime ratios?

_Points:_ 2

_Type your answer here, replacing this text._

<!-- END QUESTION -->

There are many things you can explore in this data set, but that is enough for now!

---

## 🏈 2024-25 CCSF Football Season

You are going to analyze how well the CCSF football team performed in the 2024-25 season. A football game is divided into four periods, called quarters. The number of points CCSF scored in each quarter (e.g. `CCSF 1Q`), and the number of points their opponent scored in each quarter  (e.g. `Opp 1Q`) are stored in a table called `games`. If the game is tied at the end of the 4th quarter, the game goes into an additional period called overtime.

**Notes**:
* The 2024-25 season data was collected using the Game Log from [the CCSF Athletics website](https://ccsfathletics.com/sports/fball/2024-25/teams/sanfrancisco?view=gamelog).
* The table `games_full` contains additional statistics at the right end of the table for your curiosity.

```python
games_full = Table().read_table("ccsf_fb_24_25.csv")
games = games_full.select(1, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12)
games.show()
```

Let's start by finding the total points each team scored in a game.

---

### Task 12 📍

Write a function called `sum_scores`.  It should take five arguments, where each argument is the team's quarter or overtime score. It should return the team's total score for that game.

_Points:_ 2

```python
def sum_scores(...):
    """Returns the total score calculated by adding up the score of each quarter and the overtime"""
    ...

sum_scores(14, 7, 3, 0, 3) #DO NOT CHANGE THIS LINE
```

```python
grader.check("task_12")
```

---

### Task 13 📍

Create a new table `final_scores` with three columns in this *specific* order: `Opponent`, `CCSF Score`, `Opponent Score`. You will have to create the `CCSF Score` and `Opponent Score` columns. Use the function `sum_scores` you just defined in the previous question for this problem.

*Hint:* If you want to apply a function that takes in multiple arguments, you can pass multiple column names as arguments in `tbl.apply()`. The column values will be passed into the corresponding arguments of the function. Take a look at the python reference for syntax.

*Tip:* If you’re running into issues creating final_scores, check that `ccsf_scores` and `opponent_scores` output what you want. Also, you will more than likely want to create more steps/intermediate variables.



_Points:_ 3

```python
final_scores = ...
final_scores
```

```python
grader.check("task_13")
```

We can get specific row objects from a table. You can use `tbl.row(n)` to get the `(n+1)`th row of a table. `row.item("column_name")` will allow you to select the element that corresponds to `column_name` in a particular row. 

Run the following cells to an example:

```python
games.row(9)
```

```python
games.row(9).item("CCSF 4Q")
```

---

### Task 14 📍

We want to see for a particular game whether or not CCSF won. Write a function called `did_ccsf_win`.  It should take one argument: a **row object** from the `final_scores` table. It should return either `True` if CCSF's score was greater than the Opponent's score, and `False` otherwise.

*Note 1*: "Row object" means a row from the table extracted (behind the scenes) using `tbl.row(index)` that contains all the data for that specific row. It is **not** the index of a row. Do not try and call `final_scores.row(row)` inside of the function.

*Note 2*: If you're still confused by row objects, try printing out `final_scores.row(1)` in a new cell to visually see what it looks like! This piece of code is pulling out the row object located at index 1 of the `final_scores` table and returning it. When you display it in a cell, you'll see that it is not located within a table, but is instead a standalone row object!




_Points:_ 3

```python
def ...(row):
    ...
```

```python
grader.check("task_14")
```

---

### Task 15 📍

Unfortunately, CCSF did not win against every opponent during the 2024-25 season. Using the `final_scores` table, assign `results` to an array of `True` and `False` values that correspond to whether or not CCSF won. Add the `results` array to the `final_scores` table, and assign this to `final_scores_with_results`. Then, respectively assign the number of wins and losses CCSF had to `ccsf_wins` and `ccsf_losses`.

*Hint*: When you only pass a function name and no column labels through `tbl.apply()`, the function gets applied to every row in `tbl`



_Points:_ 2

```python
results = ...
final_scores_with_results = ...
ccsf_wins = ...
ccsf_losses = ...

# Don't delete or edit the following line:
print(f"In the 2024-25 Season, CCSF Football won {ccsf_wins} games and lost {ccsf_losses} games.")
```

```python
grader.check("task_15")
```

---

### Task 16 📍

Sometimes in football, the two teams are equally matched and the game is quite close. Other times, it is a blowout, where the winning team wins by a large margin of victory. Let's define a **big win** to be a game in which the winning team won by more than 10 points. 

Create a function called `is_big_win`.
* The function should accept a single row (`datascience.tables.Row` data type) from a table like the `final_scores` table.
* The function should output `True` if CCSF won by **more than** 10 points. Otherwise, it should return `False`.

Test your function on the sixth row of the table. `final_scores.row(5)` should be the row object `Row(Opponent='San Joaquin Delta', CCSF Score=31, Opponent Score=0)`. Since CCSF's score is more than 10 points larger than San Joaquin Delta's score, then `is_big_win` applied to that row should output `True`.

_Points:_ 3

```python
...

# Test your function 
row6 = final_scores.row(5)
is_big_win(row6)
```

```python
grader.check("task_16")
```

---

### Task 17 📍

Use your `final_scores` table with your `is_big_win` function to assign `big_wins` to an array of team names that CCSF had big wins against during the 2024-25 football season.

_Points:_ 3

```python
big_wins = ...

for row_index ...
    row = ...
    opponent = ...
    if ...:
        big_wins = np.append(big_wins, ...)

big_wins
```

```python
grader.check("task_17")
```

---

## Submit Your Assignment to Canvas

Follow these steps to submit your homework assignment:

1. **Review the Rubric:** View the rubric on the associated Canvas Assignment page to understand the scoring criteria.
2. **Run the Auto-Grader:** Ensure you have executed the code cell containing the command `grader.check_all()` to run all tests for auto-graded tasks marked with 📍. This command will execute all auto-grader tests sequentially.
3. **Complete Manually Graded Tasks:** Verify that you have responded to all the manually graded tasks marked with 📍🔎.
4. **Save Your Work:** In the notebook's Toolbar, go to `File -> Save Notebook` to save your work and create a checkpoint.
5. **Download the Notebook:** In the notebook's Toolbar, go to `File -> Download IPYNB` to download the notebook (`.ipynb`) file.
6. **Upload to Canvas:** On the Canvas Assignment page, click "Start Assignment" or "New Attempt" to upload the downloaded `.ipynb` file.

---

## Attribution

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a> and derived from the <a href="https://www.data8.org/">Data 8: The Foundations of Data Science</a> offered by the University of California, Berkeley.

<img src="./by-nc-sa.png" width=100px>

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for task_01 (points: 2)
```
>>> isinstance(num_employee_ids, (int, np.integer))
True
```
failure_message: ❌ num_employee_ids should represent an integer.
success_message: ✅ num_employee_ids represents an integer.


Tests for task_03 (points: 2)
```
>>> isinstance(max_compensation, (float, np.float64, np.float32))
True
```
failure_message: ❌ max_compensation should represent a float.
success_message: ✅ max_compensation represents a float.
```
>>> isinstance(max_compensation_position, str)
True
```
failure_message: ❌ max_compensation_position should represent a string.
success_message: ✅ max_compensation_position represents a string.


Tests for task_06 (points: 3)
```
>>> isinstance(top_5_departments, Table)
True
```
failure_message: ❌ top_5_departments should be a Table.
success_message: ✅ top_5_departments is a Table.
```
>>> top_5_departments.num_rows == 5 and top_5_departments.num_columns == 2
True
```
failure_message: ❌ top_5_departments should have 5 rows and 2 columns.
success_message: ✅ top_5_departments has the correct number of rows and columns


Tests for task_07 (points: 2)
```
>>> callable(total_compensation_range)
True
```
failure_message: ❌ total_compensation_range should be callable.
success_message: ✅ total_compensation_range is callable.
```
>>> all([total_compensation_range(np.array([50000, 70000, 80000, 60000, 90000])) == 40000, total_compensation_range(np.array([60000, 60000, 60000])) == 0, total_compensation_range(np.array([75000])) == 0, total_compensation_range(np.array([-20000, -10000, 0, 10000, 20000])) == 40000])
True
```
failure_message: ❌ Your function total_compensation_range doesn't seem to be working correctly.
success_message: ✅ Your function total_compensation_range seems to be working correctly.


Tests for task_08 (points: 3)
```
>>> department_ranges.num_rows
52
```
failure_message: ❌ department_ranges doesn't seem to have the correct number of rows.
success_message: ✅ department_ranges has the correct number of rows.
```
>>> department_ranges.num_columns
7
```
failure_message: ❌ department_ranges doesn't seem to have the correct number of columns.
success_message: ✅ department_ranges has the correct number of columns.
```
>>> {'department', 'Community Health'}.issubset(set(department_ranges.labels))
True
```
failure_message: ❌ 'department' and 'Community Health' should be labels in department_ranges.
success_message: ✅ 'department' and 'Community Health' are labels in department_ranges.
```
>>> department_ranges.where('department', 'Public Health').column('Community Health').item(0)
684047.78
```
failure_message: ❌ The values in one of the columns of department_ranges don't seem correct.
success_message: ✅ The values in one of the columns of department_ranges seem correct.


Tests for task_09 (points: 3)
```
>>> isinstance(sf_above_35k, Table)
True
```
failure_message: ❌ sf_above_35k should be a Table.
success_message: ✅ sf_above_35k is a Table.
```
>>> sf_above_35k.num_columns == sf.num_columns + 1
True
```
failure_message: ❌ sf_above_35k doesn't have the correct number of columns.
success_message: ✅ sf_above_35k has the correct number of columns.


Tests for task_12 (points: 2)
```
>>> sum_scores(14, 7, 3, 0, 3) == 27 and sum_scores(2, 3, 6, 1, 0) == 12
True
```
failure_message: ❌ Your function doesn't seem to work correctly with our 4 quarter scores and over time score.
success_message: ✅ Your function works with our 4 quarter scores and over time score.


Tests for task_13 (points: 3)
```
>>> final_scores.num_columns == 3
True
```
failure_message: ❌ final_scores does not have the correct number of columns.
success_message: ✅ final_scores has the correct number of columns.
```
>>> set(['Opponent', 'CCSF Score', 'Opponent Score']) == set(final_scores.labels)
True
```
failure_message: ❌ final_scores does not have the correct labels.
success_message: ✅ final_scores has the correct labels.


Tests for task_14 (points: 3)
```
>>> 'did_ccsf_win' in globals() and callable(did_ccsf_win)
True
```
failure_message: ❌ You didn't create a function called did_ccsf_win.
success_message: ✅ You've created a function called did_ccsf_win.
```
>>> did_ccsf_win(final_scores.row(6))
True
```
failure_message: ❌ Your function doesn't that CCSF won against Laney.
success_message: ✅ Your function shows that CCSF won against Laney.
```
>>> did_ccsf_win(final_scores.row(0))
False
```
failure_message: ❌ Your function doesn't that CCSF lost against Santa Rosa.
success_message: ✅ Your function shows that CCSF lost against Santa Rosa


Tests for task_15 (points: 2)
```
>>> 0 <= ccsf_wins <= 11
True
```
failure_message: ❌ There were only 13 games, so CCSF won 0, 1, 2, ..., or 11 games.
success_message: ✅ ccsf_wins is a possible value
```
>>> 0 <= ccsf_losses <= 11
True
```
failure_message: ❌ There were only 13 games, so CCSF lost 0, 1, 2, ..., or 11 games.
success_message: ✅ ccsf_losses is a possible value


Tests for task_16 (points: 3)
```
>>> callable(is_big_win)
True
```
failure_message: ❌ It doesn't seem like you defined a function called is_big_win.
success_message: ✅ It seems like you defined a function called is_big_win.
```
>>> is_big_win(final_scores.row(0)) == True
False
```
failure_message: ❌ Your function doesn't produce the correct output for the first row of final_scores.
success_message: ✅ Your function produces the correct output for the first row of final_scores.


Tests for task_17 (points: 3)
```
>>> isinstance(big_wins, np.ndarray)
True
```
failure_message: ❌ big_wins should be a NumPy array.
success_message: ✅ big_wins is a NumPy array.
```
>>> type(big_wins.item(0)) == str
True
```
failure_message: ❌ The first item in big_wins is not a string.
success_message: ✅ The first item in big_wins is a string.
```
>>> 0 <= len(big_wins) <= 11
True
```
failure_message: ❌ There is not a possible value for the number items in big_wins.
success_message: ✅ There is a possible value for the number items in big_wins.
