```python
# Initialize Otter
import otter
grader = otter.Notebook("project2.ipynb")
```

<img src="./ccsf.png" alt="CCSF Logo" width=200px style="margin:0px -5px">

# Project 2: Final Project

In this project, you will build a classification model that guesses whether a movie is a comedy or a thriller by using the number of times certain words appear within select conversations between characters from the movies screenplay.

---

## 🎗️ Assignment Reminders

As you work on this project, consider the following:
- **Initialize Otter.** 🚨 Make sure to run the code cell at the top of this notebook that starts with `# Initialize Otter` to load the auto-grader.
- **Complete the Tasks.** Your Tasks are categorized as auto-graded (📍) and manually graded (📍🔎).
    - For all the auto-graded tasks:
        - Replace the `...` in the provided code cell with your own code.
        - Run the `grader.check` code cell to run some tests on your code.
        - Keep in mind that for homework and project assignments, sometimes there are hidden tests that you will not be able to see the results of that we use for scoring the correctness of your response. **Passing the auto-grader does not guarantee that your answer is correct.**
    - For all the manually graded tasks:
        - You might need to provide your own response to the provided prompt. Do so by replacing the template text "_Type your answer here, replacing this text._" with your own words.
        - You might need to produce a graphic or something else using code. Do so by replacing the `...` in the code cell to generate the image, table, etc.
        - In either case, [review the rubric](https://community.canvaslms.com/t5/Student-Guide/How-do-I-view-the-rubric-for-my-assignment/ta-p/275) on the associated <a href="https://ccsf.instructure.com" target="_blank">Canvas</a> Assignment page to understand the scoring criteria.
- **Code Sharing.** By submitting this project, you agree that you will not share your code directly with anybody but your partner. You are welcome to discuss questions with others but don't share the answers directly. The experience of solving the problems in this project will prepare you for our exams and potentially for future work in this field. If someone asks you for the answer, resist! Instead, you can demonstrate how you would solve a similar problem or you can focus on a specific part of a task.
- **Support.** You are not alone! Review the [Course Support page](https://ccsf-math-108.github.io/materials-sp25/resources/course-support.html) for how to get support in this course. If you're ever feeling overwhelmed or don't know how to make progress, talk to your instructor or a tutor.
- **Advice.** Develop your answers incrementally. To perform a complicated table manipulation, break it up into steps, perform each step on a different line, give a new name to each result, and check that each intermediate result is what you expect. You can add any additional names or functions you want to the provided cells. Make sure that you are using distinct and meaningful variable names throughout the notebook.
- **Variable Names.** Throughout this assignment and all future ones, please be sure to not re-assign variables throughout the notebook! _For example, if you use `max_temperature` in your answer to one question, do not reassign it later on. Otherwise, you will fail tests that you thought you were passing previously!_
- **Re-Submitting.** You may [submit](#🏁-Submit-Your-Assignment-to-Canvas) this assignment as many times as you want before the deadline. Your instructor will score the last version you submit once the deadline has passed.


---

## Configure the Notebook

Run the following cell to configure this Notebook.

```python
import datascience
import statsmodels
from datascience import *
import numpy as np
%matplotlib inline
import matplotlib.pyplot as plt
plt.style.use('fivethirtyeight')
import plotly.express as px
import warnings
warnings.simplefilter("ignore")
```

---

## 💾 Part 1: The Dataset

---

In this project, we are exploring movie screenplays. We'll be trying to predict each movie's genre from the text of its screenplay. In particular, we have compiled a list of 5,000 words that occur in conversations between movie characters. For each movie, our dataset tells us the frequency with which each of these words occurs in certain conversations in its screenplay. All words have been converted to lowercase.

Run the cell below to read the `movies` table. **It may take up to a minute to load.**

```python
movies = Table.read_table('movies.csv')
movies
```

Here is one row of the table and some of the frequencies of words that were said in the movie.

```python
movies.where("Title", "runaway bride").select(0, 1, 2, 3, 4, 14, 49, 1042, 4004)
```

The above cell prints a few columns of the row for the comedy movie *Runaway Bride*.  The conversations examined in the movie contain 4895 words. The word "it" appears 115 times, as it makes up  $\frac{115}{4895} \approx 0.0234092$ of the words in the conversations. The word "england" doesn't appear at all.

Additional context: This numerical representation of a body of text, one that describes only the frequencies of individual words, is called a bag-of-words representation. This is a model that is often used in [NLP](https://en.wikipedia.org/wiki/Natural_language_processing). A lot of information is discarded in this representation: the order of the words, the context of each word, who said what, the cast of characters and actors, etc. However, a bag-of-words representation is often used for machine learning applications as a reasonable starting point, because a great deal of information is also retained and expressed in a convenient and compact format. 

In this project, we will investigate whether this representation is sufficient to build an accurate genre classifier.

All movie titles are unique. The `row_for_title` function provides fast access to the one row for each title. 

*Note: All movies in our dataset have their titles lower-cased.* 

```python
title_index = movies.index_by('Title')
def row_for_title(title):
    """Return the row for a title, similar to the following expression (but faster)
    
    movies.where('Title', title).row(0)
    """
    return title_index.get(title)[0]

row_for_title('toy story');
```

For example, the fastest way to find the frequency of "fun" in the movie *Toy Story* is to access the `'fun'` item from its row. Check the original table to see if this worked for you!

```python
row_for_title('toy story').item('fun') 
```

---

#### Task 1.1 📍

Set `expected_row_sum` to the number that you __expect__ will result from summing all proportions in each row, excluding the first five columns. Think about what any one row adds up to.

_Points:_ 2

```python
# Set row_sum to a number that's the (approximate) sum of each row of word proportions.
expected_row_sum = ...
```

```python
grader.check("task_1_1")
```

---

This dataset was extracted from [a dataset from Cornell University](http://www.cs.cornell.edu/~cristian/Cornell_Movie-Dialogs_Corpus.html). After transforming the dataset (e.g., converting the words to lowercase, removing the naughty words, and converting the counts to frequencies), we created this new dataset containing the frequency of 5,000 common words in each movie.

```python
print('Words with frequencies:', movies.drop(np.arange(5)).num_columns) 
print('Movies with genres:', movies.num_rows)
```

---

### Section 1.1: Word Stemming

The columns other than "Title", "Year", "Rating", "Genre", and "# Words" in the `movies` table are all words that appear in some of the conversations in the movies within our dataset.  These words have been *stemmed*, or abbreviated heuristically, in an attempt to make different [inflected](https://en.wikipedia.org/wiki/Inflection) forms of the same base word into the same string.  For example, the column "manag" is the sum of proportions of the words "manage", "manager", "managed", and "managerial" (and perhaps others) in each movie. This is a common technique used in machine learning and natural language processing.

Stemming makes it a little tricky to search for the words you want to use, so we have provided another table called `vocab_table` that will let you see examples of unstemmed versions of each stemmed word. Run the code below to load it.

**Note:** You should use `vocab_table` for the rest of Section 1.1, not `vocab_mapping`. 

```python
# Just run this cell.
vocab_mapping = Table.read_table('stem.csv')
stemmed = np.take(movies.labels, np.arange(3, len(movies.labels)))
vocab_table = Table().with_column('Stem', stemmed).join('Stem', vocab_mapping)
vocab_table.take(np.arange(1100, 1110))
```

---

#### Task 1.1.1 📍

Using `vocab_table`, find the stemmed version of the word "elements" and assign the value to `stemmed_message`.


_Points:_ 3

```python
stemmed_message = ...
stemmed_message
```

```python
grader.check("task_1_1_1")
```

---

#### Task 1.1.2 📍

What stem in the dataset has the most words that are associated to it? Assign `most_stem` to that stem. 

_Points:_ 2

```python
most_stem = ...
most_stem
```

```python
grader.check("task_1_1_2")
```

---

#### Task 1.1.3 📍

What is the longest word in the dataset that wasn't shortened to its stem? Assign that to `longest_uncut`. Break ties alphabetically from Z to A (so if your options are "cat" or "bat", you should pick "cat"). Note that when sorting letters, the letter `a` is smaller than the letter `z`.

*Hint 1:* `vocab_table` has 2 columns: one for stems and one for the unstemmed (normal) word. Find the longest word that wasn't cut at all (same length as stem).

*Hint 2:* There is a table function that allows you to compute a function on every element in a column. Review the [Quick Reference](https://ccsf-math-108.github.io/materials-fa24/resources/quick-reference.html) if you aren't sure which one.

*Hint 3:* In our solution, we found it useful to first add columns with the length of the word and the length of the stem, and then to add a column with the difference between those lengths. What will the difference be if the word is not shortened?

*Hint 4:* You may also want to sort more than once.

_Points:_ 2

```python
tbl_with_lens = ...
tbl_with_diff = ...

longest_uncut = ...
longest_uncut
```

```python
grader.check("task_1_1_3")
```

---

#### Task 1.1.4 📍

How many stems have only one word that is shortened to them? For example, if the stem "book" only maps to the word "books" and if the stem "a" only maps to the word "a," both should be counted as stems that map only to a single word.

Assign `count_single_stems` to the count of stems that map to one word only. 

_Points:_ 2

```python
count_single_stems = ...
count_single_stems
```

```python
grader.check("task_1_1_4")
```

---

### Section 1.2: Linear Regression

Let's explore our dataset before trying to build a classifier. To start, we’ll use the associated proportions to investigate the relationship between different words.

The first association we'll investigate is the association between the proportion of words that are "outer" and the proportion of words that are "space". 

As usual, we'll investigate our data visually before performing any numerical analysis.

Run the cell below to plot a scatter diagram of "space" proportions vs "outer" proportions and to create the `outer_space` table. Each point on the scatter plot represents one movie.

```python
# Just run this cell!
outer_space = movies.select("outer", "space")
outer_space.scatter("outer", "space")
plt.axis([-0.0005, 0.001, -0.0005, 0.003]);
plt.xticks(rotation=45);
```

---

#### Task 1.2.1 📍

Looking at that chart it is difficult to see if there is an association. Calculate the correlation coefficient for the potential linear association between proportion of words that are "outer" and the proportion of words that are "space" for every movie in the dataset, and assign it to `outer_space_r`. 

*Hint:* If you need a refresher on how to calculate the correlation coefficient check out [Ch 15.1](https://inferentialthinking.com/chapters/15/1/Correlation.html#calculating-r).

_Points:_ 2

```python
# These two arrays should make your code cleaner!
outer = movies.column("outer")
space = movies.column("space")

outer_su = ...
space_su = ...

outer_space_r = ...
outer_space_r
```

```python
grader.check("task_1_2_1")
```

---

#### Task 1.2.2 📍🔎

<!-- BEGIN QUESTION -->

Choose two *different* words in the dataset with a magnitude (absolute value) of correlation higher than 0.2 and plot a scatter plot with a line of best fit for them. **Do not pick "outer" and "space" or "san" and "francisco".** The code to plot the scatter plot and line of best fit is given for you, you just need to calculate the correct values to `r`, `slope` and `intercept`. 

*Hint 1:* It's easier to think of words with a positive correlation, i.e. words that are often mentioned together*. Try to think of common phrases or idioms.

*Hint 2:* Refer to [Section 15.2](https://inferentialthinking.com/chapters/15/2/Regression_Line.html#units-of-measurement-of-the-slope) of the textbook for the formulas. 

_Points:_ 2

```python
word_x = ...
word_y = ...

# These arrays should make your code cleaner!
arr_x = movies.column(word_x)
arr_y = movies.column(word_y)

x_su = ...
y_su = ...

r = ...

slope = ...
intercept = ...

# DON'T CHANGE THESE LINES OF CODE
movies.scatter(word_x, word_y)
max_x = max(movies.column(word_x))
plt.title(f"Correlation: {r}, magnitude greater than .2: {abs(r) >= 0.2}")
plt.plot([0, max_x * 1.3], [intercept, intercept + slope * (max_x*1.3)], color='gold');
```

<!-- END QUESTION -->

---

#### Task 1.2.3 📍

Imagine that you picked the words "san" and "francisco" as the two words that you would expect to be correlated because they compose the city name San Francisco. Assign `san_francisco` to either the number 1 or the number 2 according to which statement is true regarding the correlation between "san" and "francisco." 

1. "san" can also preceed other city names like San Diego and San Jose. This might lead to "san" appearing in movies without "francisco," and would reduce the correlation between "san" and "francisco."
2. "san" can also preceed other city names like San Diego and San Jose. The fact that "san" could appear more often in front of different cities and without "francisco" would increase the correlation between "san" and "francisco." 

_Points:_ 2

```python
san_francisco = ...
```

```python
grader.check("task_1_2_3")
```

---

### Section 1.3: Splitting the dataset

Now, we're going to use our `movies` dataset for two purposes.

1. First, we want to *train* movie genre classifiers.
2. Second, we want to *test* the performance of our classifiers. 

Hence, we need two different datasets: *training* and *test*. 

The purpose of a classifier is to classify unseen data that is similar to the training data. The test dataset will help us determine the accuracy of our predictions by comparing the actual genres of the movies with the genres that our classifier predicts. Therefore, we must ensure that there are no movies that appear in both sets. We do so by splitting the dataset randomly. The dataset has already been permuted randomly, so it's easy to split.  We just take the first 85% of the dataset for training and the rest for test. 

Run the code below (without changing it) to separate the datasets into two tables.

```python
# Here we have defined the proportion of our data
# that we want to designate for training as 17/20ths
# of our total dataset.  3/20ths of the data is
# reserved for testing.

training_proportion = 17/20

num_movies = movies.num_rows
num_train = int(num_movies * training_proportion)
num_test = num_movies - num_train

train_movies = movies.take(np.arange(num_train))
test_movies = movies.take(np.arange(num_train, num_movies))

print("Training: ",   train_movies.num_rows, ";",
      "Test: ",       test_movies.num_rows)
```

---

#### Task 1.3.1 📍🔎

<!-- BEGIN QUESTION -->

Draw a horizontal bar chart with two bars that show the proportion of Comedy movies in each dataset (`train_movies` and `test_movies`). The two bars should be labeled "Training" and "Test". Complete the function `comedy_proportion` first; it should help you create the bar chart. 

*Hint*: Refer to [Section 7.1](https://inferentialthinking.com/chapters/07/1/Visualizing_Categorical_Distributions.html#bar-chart) of the textbook if you need a refresher on bar charts.

_Points:_ 2

```python
def comedy_proportion(table):
    # Return the proportion of movies in a table that have the comedy genre.
    ...
    return ...

# The staff solution took multiple lines.  Start by creating a table.
# If you get stuck, think about what sort of table you need for barh to work
...
```

<!-- END QUESTION -->

---

## 🏘️ Part 2: K-Nearest Neighbors - A Guided Example

---

[K-Nearest Neighbors (k-NN)](https://inferentialthinking.com/chapters/17/1/Nearest_Neighbors.html) is a classification algorithm.  Given some numerical *attributes* (also called *features*) of an unseen example, it decides which category that example belongs to based on its similarity to previously seen examples. Predicting the category of an example is called *labeling*, and the predicted category is also called a *label*.

An attribute (feature) we have about each movie is *the proportion of times a particular word appears in the selected conversations from the movie*, and the labels are two movie genres: comedy and thriller.  The algorithm requires many previously seen examples for which both the attributes and labels are known: that's the `train_movies` table.

To build understanding, we're going to visualize the algorithm instead of just describing it.

---

### Section 2.1: Classifying a movie

In k-NN, we classify a movie by finding the `k` movies in the *training set* that are most similar according to the features we choose. We call those movies with similar features the *nearest neighbors*.  The k-NN algorithm assigns the movie to the most common category among its `k` nearest neighbors.

Let's limit ourselves to just 2 features for now, so we can plot each movie.  The features we will use are the proportions of the words "water" and "feel" in the conversations.  Taking the movie *Monty Python and the Holy Grail* (in the test set), 0.000804074 of its words are "water" and 0.0010721 are "feel". This movie appears in the test set, so let's imagine that we don't yet know its genre.

First, we need to make our notion of similarity more precise.  We will say that the *distance* between two movies is the straight-line distance between them when we plot their features on a scatter diagram. 

**This distance is called the Euclidean ("yoo-KLID-ee-un") distance, whose formula is $\sqrt{(x_1 - x_2)^2 + (y_1 - y_2)^2}$.**

For example, in the movie *Clerks.* (in the training set), 0.00016293 of all the words in the movie conversations are "water" and 0.00154786 are "feel".  Its distance from *Monty Python and the Holy Grail* on this 2-word feature set is $\sqrt{(0.000804074 - 0.000162933)^2 + (0.0010721 - 0.00154786)^2} \approx 0.000798379$.  (If we included more or different features, the distance could be different.)

A third movie, *The Godfather* (in the training set), has 0 "water" and 0.00015122 "feel". 

The function below creates a plot to display the "water" and "feel" features of a test movie and some training movies. As you can see in the result, *Monty Python and the Holy Grail* is more similar to *Clerks.* than to the *The Godfather* based on these features, which makes sense as both movies are comedy movies, while *The Godfather* is a thriller.

```python
# Just run this cell.
def plot_with_two_features(test_movie, training_movies, x_feature, y_feature):
    """Plot a test movie and training movies using two features."""
    test_row = row_for_title(test_movie)
    distances = Table().with_columns(
            x_feature, [test_row.item(x_feature)],
            y_feature, [test_row.item(y_feature)],
            'Color',   ['unknown'],
            'Title',   [test_movie]
        )
    for movie in training_movies:
        row = row_for_title(movie)
        distances.append([row.item(x_feature), row.item(y_feature), row.item('Genre'), movie])
    distances.scatter(x_feature, y_feature, group='Color', labels='Title', s=50)
    
training = ["clerks.", "the godfather"] 
plot_with_two_features("monty python and the holy grail", training, "water", "feel")
plt.axis([-0.0008, 0.001, -0.004, 0.007]);
```

---

#### Task 2.1.1 📍

Compute the Euclidean distance (defined in the section above) between the two movies, *Monty Python and the Holy Grail* and *The Godfather*, using the `water` and `feel` features only.  Assign it the name `one_distance`. 

*Hint 1:* If you have a row, you can use `item` to get a value from a column by its name.  For example, if `r` is a row, then `r.item("Genre")` is the value in column `"Genre"` in row `r`.

*Hint 2:* Refer to the beginning of Part 1 if you don't remember what `row_for_title` does.

*Hint 3:* In the formula for Euclidean distance, think carefully about what `x` and `y` represent. Refer to the example in the text above if you are unsure.

_Points:_ 2

```python
python = row_for_title("monty python and the holy grail") 
godfather = row_for_title("the godfather") 

one_distance = ...
one_distance
```

```python
grader.check("task_2_1_1")
```

---

Below, we've added a third training movie, *The Silence of the Lambs*. Before, the point closest to *Monty Python and the Holy Grail* was *Clerks.*, a comedy movie. However, now the closest point is *The Silence of the Lambs*, a thriller movie.

```python
training = ["clerks.", "the godfather", "the silence of the lambs"] 
plot_with_two_features("monty python and the holy grail", training, "water", "feel") 
plt.axis([-0.0008, 0.001, -0.004, 0.007]);
```

---

#### Task 2.1.2 📍

Complete the function `distance_two_features` that computes the Euclidean distance between any two movies, using two features. The last two lines call your function to show that *Monty Python and the Holy Grail* is closer to *The Silence of the Lambs* than it is to *Clerks*. 

_Points:_ 2

```python
def distance_two_features(title0, title1, x_feature, y_feature):
    """Compute the distance between two movies with titles title0 and title1.
    
    Only the features named x_feature and y_feature are used when computing the distance.
    """
    row0 = ...
    row1 = ...
    ...

for movie in make_array("clerks.", "the silence of the lambs"):
    movie_distance = distance_two_features(movie, "monty python and the holy grail", "water", "feel")
    print(movie, 'distance:\t', movie_distance)
```

```python
grader.check("task_2_1_2")
```

---

#### Task 2.1.3 📍

Define the function `distance_from_python` so that it works as described in its documentation. 

**Note:** Your solution should not use arithmetic operations directly. Instead, it should make use of existing functionality above!

_Points:_ 2

```python
def distance_from_python(title):
    """The distance between the given movie and "monty python and the holy grail", 
    based on the features "water" and "feel".
    
    This function takes a single argument:
      title: A string, the name of a movie.
    """
    
    ...

# Calculate the distance between "Clerks." and "Monty Python and the Holy Grail"
distance_from_python('clerks.')
```

```python
grader.check("task_2_1_3")
```

---

#### Task 2.1.4 📍

Using the features `"water"` and `"feel"`, what are the names and genres of the 5 movies in the **training set** closest to *Monty Python and the Holy Grail*?  To address this task, make a **table** named `close_movies` containing those 5 movies with columns `"Title"`, `"Genre"`, `"water"`, and `"feel"`, as well as a column called `"distance from python"` that contains the distance from *Monty Python and the Holy Grail*.  The table should be **sorted in ascending order by `"distance from python"`**.

*Note:* Why are smaller distances from *Monty Python and the Holy Grail* more helpful in helping us classify the movie?

*Hint:* Your final table should only have 5 rows. How can you get the first five rows of a table?

_Points:_ 4

```python

# The staff solution took multiple lines.
...
close_movies = ...
close_movies
```

```python
grader.check("task_2_1_4")
```

---

#### Task 2.1.5 📍

Next, we'll clasify *Monty Python and the Holy Grail* based on the genres of the closest movies. 

To do so, define the function `most_common` so that it works as described in its documentation below. 

_Points:_ 1

```python
def most_common(label, table):
    """The most common element in a column of a table.
    
    This function takes two arguments:
      label: The label of a column, a string.
      table: A table.
     
    It returns the most common value in the label column of the table.
    In case of a tie, it returns any one of the most common values.    
    """
    ...

# Calling most_common on your table of 5 nearest neighbors classifies
# "monty python and the holy grail" as a thriller movie, 3 votes to 2. 
most_common('Genre', close_movies)
```

```python
grader.check("task_2_1_5")
```

---

Congratulations are in order -- you've classified your first movie! However, we can see that the classifier doesn't work too well since it categorized *Monty Python and the Holy Grail* as a thriller movie (unless you count the thrilling holy hand grenade scene). Let's see if we can do better!

---

## ✨ Part 3: Features

---

Now, we're going to extend our classifier to consider more than two features at a time.

Euclidean distance still makes sense with more than two features. For `n` different features, we compute the difference between corresponding feature values for two movies, square each of the `n`  differences, sum up the resulting numbers, and take the square root of the sum.

---

### Task 3.1 📍

Write a function called `distance` to compute the Euclidean distance between two **arrays** of **numerical** features (e.g. arrays of the proportions of times that different words appear). The function should be able to calculate the Euclidean distance between two arrays of arbitrary (but equal) length.

Next, use the function you just defined to compute the distance **between the first and second movie** in the **training set** *using all of the features*.  (Remember that the first five columns of your tables are not features.)

*Hint 1:* To convert rows to arrays, use `np.array`. For example, if `t` was a table, `np.array(t.row(0))` converts row 0 of `t` into an array. We've included two optional variables `array_values_movie_1` and `array_values_movie_2` to help you focus on creating two feature arrays to test your function.

*Hint 2:* Make sure to drop the first five columns of the table before you compute `distance_first_to_second`, as these columns do not contain any features (the proportions of words). 

_Points:_ 7

```python
def distance(features_array1, features_array2):
    """The Euclidean distance between two arrays of feature values."""
    ...

array_values_movie_1 = ...
array_values_movie_2 = ...
distance_first_to_second = ...
distance_first_to_second
```

```python
grader.check("task_3_1")
```

---

### Section 3.1: Creating your own feature set

Unfortunately, using all of the features has some downsides.  One clear downside is the lack of *computational efficiency* -- computing Euclidean distances just takes a long time when we have lots of features.  You might have noticed that in the last task!

So, we are going to have you select just 20.  We'd like to choose features that are very *discriminative*, that is, features that lead  the classifier to correctly classify as much of the test set as possible.  This process of choosing features that will make a classifier work well is sometimes called *feature selection*, or, more broadly, *feature engineering*.

To begin your feature selection, run the following code to create a scatter plot. The scatter plot shows each word as a dot and the position of each dot is determined by the average occurrence of that word within comedy movies (horizontal axis) and the average occurrence of that word within thriller movies (vertical axis). 

*Note 1:* The line graphed is the linear regression line, not the line $y = x$.

*Note 2*: The interactive graphing library Plotly Express was used to create the scatter plot. When you have over the graph, a menu bar will appear that allows you to zoom in/out on various parts of the graph using the navigation bar and hover over points to see the word each point represents.

```python
# Run this cell to generate the interactive image
words_with_averages = Table().read_table('word_averages.csv')

words_with_averages_df = words_with_averages.to_df()
fig = px.scatter(words_with_averages_df, 
                 x="Average Occurrence (Comedy)", 
                 y="Average Occurrence (Thriller)", 
                 hover_data=['Word'], 
                 trendline="ols",
                 trendline_color_override="pink",
                 title="Original Units")
fig.update_layout(autosize=False, width=600, height=600)
fig.show()
```

Tasks 3.1.1 through 3.1.4 will ask you to interpret the plot above. For each task, select one of the following choices and assign its number to the provided name.
1. The word is common in both comedy and thriller movies 
2. The word is uncommon in comedy movies and common in thriller movies
3. The word is common in comedy movies and uncommon in thriller movies
4. The word is uncommon in both comedy and thriller movies
5. It is not possible to say from the plot 

---

#### Task 3.1.1 📍

What properties do words in the bottom left corner of the plot have? Your answer should be a single integer from 1 to 5, corresponding to the correct statement from the choices above.

_Points:_ 1

```python
bottom_left = ...
```

```python
grader.check("task_3_1_1")
```

---

#### Task 3.1.2 📍

What properties do words in the bottom right corner have?

_Points:_ 1

```python
bottom_right = ...
```

```python
grader.check("task_3_1_2")
```

---

#### Task 3.1.3 📍

What properties do words in the top right corner have? 

_Points:_ 1

```python
top_right = ...
```

```python
grader.check("task_3_1_3")
```

---

#### Task 3.1.4 📍

What properties do words in the top left corner have?

_Points:_ 1

```python
top_left = ...
```

```python
grader.check("task_3_1_4")
```

---

#### Task 3.1.5 📍

If we see a movie with a lot of words that are common for comedy movies but uncommon for thriller movies, what would be a reasonable guess about the genre of the movie? Assign `movie_genre` to the integer corresponding to your answer:
1. It is a thriller movie.
2. It is a comedy movie.

_Points:_ 1

```python
movie_genre_guess = ...
```

```python
grader.check("task_3_1_5")
```

---

In the next task, we want you to pick some words to help your classifier distinguish between a comedy and a thriller. Run the following cell to create two scatter plots. 

1. The first scatter is the same plot from above.
2. The second plot shows the same information but it uses a log scale so it is a little easier to distinguish between the words (points) on the scatter plot in the lower left-hand corner.

```python
fig = px.scatter(words_with_averages_df, 
                 x="Average Occurrence (Comedy)", 
                 y="Average Occurrence (Thriller)", 
                 hover_data=['Word'], 
                 trendline="ols",
                 trendline_color_override="pink",
                 title="Original Units")
fig.update_layout(autosize=False, width=600, height=600)
fig.show()

fig_log = px.scatter(words_with_averages_df, 
                     x="Average Occurrence (Comedy)", 
                     y="Average Occurrence (Thriller)", 
                     hover_data=['Word'], 
                     trendline="ols",
                     trendline_color_override="pink",
                     title="Log Scale")

fig_log.update_xaxes(type='log')
fig_log.update_yaxes(type='log')
fig_log.update_layout(autosize=False, width=600, height=600)
fig_log.show()
```

---

#### Task 3.1.6 📍

Using the plots above, make an array of at least 10 common words that you think might let you **distinguish** between comedy and thriller movies. Make sure to choose words that are **frequent enough** that every movie contains at least one of them. Don't just choose the most frequent words though--you can do much better.

_Points:_ 5

```python
# Set my_features to an array of at least 10 features (strings that are column labels)

my_features = ...

# Select the 10 features of interest from both the train and test sets
train_my_features = train_movies.select(my_features)
test_my_features = test_movies.select(my_features)
```

```python
grader.check("task_3_1_6")
```

---

This test makes sure that you have chosen words such that at least one appears in each movie. If you can't find words that satisfy this test just through intuition, try writing code to print out the titles of movies that do not contain any words from your list, then look at the words they do contain.

---

Next, let's classify the first movie from our test set using these features.  You can examine the movie by running the cells below. Do you think it will be classified correctly?

```python
print("Movie:")
test_movies.take(0).select('Title', 'Genre').show()
print("Features:")
test_my_features.take(0).show()
```

As before, we want to look for the movies in the training set that are most like our test movie.  We will calculate the Euclidean distances from the test movie (using `my_features`) to all movies in the training set.  You could do this with a `for` loop, but to make it computationally faster, we have provided a function, `fast_distances`, to do this for you.  Read its documentation to make sure you understand what it does.  (You don't need to understand the code in its body unless you want to.)

```python
# Just run this cell to define fast_distances.

def fast_distances(test_row, train_table):
    """Return an array of the distances between test_row and each row in train_table.

    Takes 2 arguments:
      test_row: A row of a table containing features of one
        test movie (e.g., test_my_features.row(0)).
      train_table: A table of features (for example, the whole
        table train_my_features)."""
    assert train_table.num_columns < 50, "Make sure you're not using all the features of the movies table."
    assert type(test_row) != datascience.tables.Table, "Make sure you are passing in a row object to fast_distances."
    assert len(test_row) == len(train_table.row(0)), "Make sure the length of test row is the same as the length of a row in train_table."
    counts_matrix = np.asmatrix(train_table.columns).transpose()
    diff = np.tile(np.array(list(test_row)), [counts_matrix.shape[0], 1]) - counts_matrix
    np.random.seed(0) # For tie breaking purposes
    distances = np.squeeze(np.asarray(np.sqrt(np.square(diff).sum(1))))
    eps = np.random.uniform(size=distances.shape)*1e-10 #Noise for tie break
    distances = distances + eps
    return distances
```

---

#### Task 3.1.7 📍

Use the `fast_distances` function provided above to compute the distance from the first movie in your test set to all the movies in your training set, **using your set of features**. Make a new table called `genre_and_distances` with one row for each movie in the training set and two columns:
* The `"Genre"` of the training movie
* The `"Distance"` from the first movie in the test set 

Ensure that `genre_and_distances` is **sorted in ascending order by distance to the first test movie**.

*Hint:* Think about how you can use the variables you defined in Task 3.1.6.

_Points:_ 4

```python
# The staff solution took multiple lines of code.
genre_and_distances = ...
genre_and_distances
```

```python
grader.check("task_3_1_7")
```

---

#### Task 3.1.8 📍

Now compute the 7-nearest neighbors classification of the first movie in the test set.  That is, decide on its genre by finding the most common genre among its 7 nearest neighbors in the training set, according to the distances you've calculated.  Then check whether your classifier chose the right genre.  (Depending on the features you chose, your classifier might not get this movie right, and that's okay.)

*Hint 1:* You should use the `most_common` function that was defined earlier.

*Hint 2:* You should use a comparison operator.

_Points:_ 2

```python
# Set my_assigned_genre to the most common genre among these.
my_assigned_genre = ...

# Set my_assigned_genre_was_correct to True if my_assigned_genre
# matches the actual genre of the first movie in the test set, False otherwise.
my_assigned_genre_was_correct = ...

print("The assigned genre, {}, was{}correct.".format(my_assigned_genre, " " if my_assigned_genre_was_correct else " not "))
```

```python
grader.check("task_3_1_8")
```

---

### Section 3.2: A classifier function


Now we can write a single function that encapsulates the whole process of classification.

---

#### Task 3.2.1 📍

Write a function called `classify`.  It should take the following four arguments:
* A row of features for a movie to classify (e.g., `test_my_features.row(0)`).
* A table with a column for each feature (e.g., `train_my_features`).
* An array of classes (e.g. the labels "comedy" or "thriller") that has as many items as the previous table has rows, and in the same order. *Hint:* What are the labels of each row in the training set? 
* `k`, the number of neighbors to use in classification.

It should return the class (the string `'comedy'` or the string `'thriller'`) a k-nearest neighbor classifier picks for the given row of features.

_Points:_ 2

```python
def classify(test_row, train_features, train_labels, k):
    """Return the most common class among k nearest neigbors to test_row."""
    distances = fast_distances(test_row, train_features)
    genre_and_distances = ...
    ...
```

```python
grader.check("task_3_2_1")
```

---

#### Task 3.2.2 📍

Assign `godzilla_genre` to the genre predicted by your classifier for the movie "godzilla" in the test set, using **15 neighbors** and using your 10 features. 

*Hint:* The `row_for_title` function will not work here.

_Points:_ 1

```python
# The staff solution first defined a row called godzilla_features.
godzilla_features = ...
godzilla_genre = ...
godzilla_genre
```

```python
grader.check("task_3_2_2")
```

---

Finally, when we evaluate our classifier, it will be useful to have a classification function that is specialized to use a fixed training set and a fixed value of `k`.

---

#### Task 3.2.3 📍

Create a classification function that takes as its argument a row containing your 10 features and classifies that row using the 15-nearest neighbors algorithm with `train_my_features` as its training set. 

_Points:_ 1

```python
def classify_feature_row(row):
    ...

# When you're done, this should produce 'thriller' or 'comedy'.
classify_feature_row(test_my_features.row(0))
```

```python
grader.check("task_3_2_3")
```

---

### Section 3.3: Evaluating your classifier

Now that it's easy to use the classifier, let's see how accurate it is on the whole test set.

---

#### Task 3.3.1 📍

Use `classify_feature_row` and `apply` to classify every movie in the test set.  Assign these guesses as an array to `test_guesses`.  Finally, calculate the accuracy of your classifier. 

*Hint 1*: If you do not specify any columns in `tbl.apply(...)`, your function will be applied to every row object in `tbl`.

*Hint 2*: Which dataset do you want to apply this function to?

_Points:_ 2

```python
test_guesses = ...
accuracy = ...
accuracy
```

```python
grader.check("task_3_3_1")
```

---

#### Task 3.3.2 📍

An important part of evaluating your classifiers is figuring out where they make mistakes. Assign the name `test_movie_correctness` to a table with three columns, `'Title'`, `'Genre'`, and `'Was correct'`. 

- The `'Genre'` column should contain the original genres, not the ones you predicted. 
- The `'Was correct'` column should contain `True` or `False` depending on whether or not the movie was classified correctly.


_Points:_ 3

```python
# Feel free to use multiple lines of code
# but make sure to assign test_movie_correctness to the proper table!
test_movie_correctness = ...
test_movie_correctness.sort('Was correct', descending = True).show(5)
```

```python
grader.check("task_3_3_2")
```

---

### Reflecting on the Mistakes

You should notice that your classifier misclassifies one genre more than the other. Why did that happen? There are many possible reasons, but one might have to do with the fact that some movies are not purely a comedy or thriller, they are some blend of genres. Consider the movie [Scary Movie](https://en.wikipedia.org/wiki/Scary_Movie)!

---

## 🎉 Reflection

Great work! You've gone through one iteration of classifier design.  Let's summarize the steps:
1. From available data, select test and training sets.
2. Choose an algorithm you're going to use for k-NN classification.
3. Identify some features.
4. Define a classifier function using your features and the training set.
5. Evaluate its performance (the proportion of correct classifications) on the test set.
Machine learning projects rarely end after a single iteration. Instead, think of this process as one that involves re-evaluating your project goals, the quality and availability of your data, and your approach to classification.

---

## 🏁 Submit Your Assignment to Canvas

Follow these steps to submit your project assignment:

1. **Review the Rubric:** View the rubric on the associated Canvas Assignment page to understand the scoring criteria.
2. **Run the Auto-Grader:** Ensure you have executed the code cell containing the command `grader.check_all()` to run all tests for auto-graded tasks marked with 📍. This command will execute all auto-grader tests sequentially.
3. **Complete Manually Graded Tasks:** Verify that you have responded to all the manually graded tasks marked with 📍🔎.
4. **Save Your Work:** In the notebook's Toolbar, go to `File -> Save Notebook` to save your work and create a checkpoint.
5. **Download the Notebook:** In the notebook's Toolbar, go to `File -> Download IPYNB` to download the notebook (`.ipynb`) file.
6. **Upload to Canvas:** On the Canvas Assignment page, click "Start Assignment" or "New Attempt" to upload the downloaded `.ipynb` file.

---

## Attribution

This content is licensed under the <a href="https://creativecommons.org/licenses/by-nc-sa/4.0/">Creative Commons Attribution-NonCommercial-ShareAlike 4.0 International License (CC BY-NC-SA 4.0)</a> and derived from the <a href="https://www.data8.org/">Data 8: The Foundations of Data Science</a> offered by the University of California, Berkeley.

<img src="./by-nc-sa.png" width=100px>

---

To double-check your work, the cell below will rerun all of the autograder tests.

```python
grader.check_all()
```

Tests for task_1_1 (points: 2)
```
>>> 0 < expected_row_sum
True
```
failure_message: ❌ expected_row_sum should be greater than 0.
success_message: ✅ expected_row_sum is greater than 0.


Tests for task_1_1_1 (points: 3)
```
>>> type(stemmed_message) in set([np.str_, str])
True
```
failure_message: ❌ stemmed_message should be a string.
success_message: ✅ stemmed_message is a string.
```
>>> len(stemmed_message) < len('elements')
True
```
failure_message: ❌ The length of stemmed_message should be less than the lenght of 'elements'.
success_message: ✅ The length of stemmed_message is less than the lenght of 'elements'.


Tests for task_1_1_2 (points: 2)
```
>>> type(most_stem) in set([str, np.str_])
True
```
failure_message: ❌ most_stem should be a string.
success_message: ✅ most_stem is a string.


Tests for task_1_1_3 (points: 2)
```
>>> type(longest_uncut) in set([str, np.str_])
True
```
failure_message: ❌ longest_uncut should be a string.
success_message: ✅ longest_uncut is a string.


Tests for task_1_1_4 (points: 2)
```
>>> type(count_single_stems) in set([int, np.int32, np.int64])
True
```
failure_message: ❌ count_single_stems should be an integer.
success_message: ✅ count_single_stems is an integer.


Tests for task_1_2_1 (points: 2)
```
>>> 0.2 < outer_space_r < 0.4
True
```
failure_message: ❌ outer_space_r should be between 0.2 and 0.4.
success_message: ✅ outer_space_r is between 0.2 and 0.4.


Tests for task_1_2_3 (points: 2)
```
>>> type(san_francisco) in set([int, np.int32, np.int64])
True
```
failure_message: ❌ san_francisco should be an integer.
success_message: ✅ san_francisco is an integer.


Tests for task_2_1_1 (points: 2)
```
>>> 0 < one_distance < 0.01
True
```
failure_message: ❌ one_distance should be between 0 and 0.01.
success_message: ✅ one_distance is between 0 and 0.01.


Tests for task_2_1_2 (points: 2)
```
>>> correct_dis = 0.001406116
>>> dis = distance_two_features('clerks.', 'the godfather', 'water', 'feel')
>>> np.isclose(np.round(dis, 9), correct_dis)
True
```
failure_message: ❌ distance_two_features('clerks.', 'the godfather', 'water', 'feel') didn't return the correct distance.
success_message: ✅ distance_two_features('clerks.', 'the godfather', 'water', 'feel') returned the correct distance.
```
>>> correct_dis = 0.005647119
>>> dis = distance_two_features('clerks.', 'the godfather', 'your', 'that')
>>> np.isclose(np.round(dis, 9), correct_dis)
True
```
failure_message: ❌ distance_two_features('clerks.', 'the godfather', 'your', 'that') didn't return the correct distance.
success_message: ✅ distance_two_features('clerks.', 'the godfather', 'your', 'that') returned the correct distance.


Tests for task_2_1_3 (points: 2)
```
>>> np.isclose(distance_from_python('clerks.'), 0.00079838)
True
```
failure_message: ❌ distance_from_python('clerks.') didn't return a reasonably correct distance.
success_message: ✅ distance_from_python('clerks.') returned a reasonably correct distance.


Tests for task_2_1_4 (points: 4)
```
>>> set(close_movies.labels) >= {'Genre', 'Title', 'feel', 'water'}
True
```
failure_message: ❌ close_movies doesn't have the correct labels.
success_message: ✅ close_movies has the correct labels.
```
>>> close_movies.num_rows == 5
True
```
failure_message: ❌ close_movies doesn't have the correct number of rows.
success_message: ✅ close_movies has the correct number of rows.
```
>>> close_movies.column('Title').item(0) != 'monty python and the holy grail'
True
```
failure_message: ❌ It seems like you are not using the training set.
success_message: ✅ It seems like you are using the training set.


Tests for task_2_1_5 (points: 1)
```
>>> [most_common('Genre', close_movies.take(range(k))) for k in range(1, 6, 2)]
['thriller', 'thriller', 'thriller']
```
failure_message: ❌ most_common didn't work correctly for our test case.
success_message: ✅ most_common worked correctly for our test case.


Tests for task_3_1 (points: 7)
```
>>> 0.0 <= distance_first_to_second <= 0.1
True
```
failure_message: ❌ distance_first_to_second should be between 0.0 and 0.1, inclusive.
success_message: ✅ distance_first_to_second is betwen 0.0 and 0.1, inclusive.
```
>>> np.isclose(distance(make_array(1, 2), make_array(1, 2)), 0)
True
```
failure_message: ❌ distance doesn't correctly work when the inputs are the same array.
success_message: ✅ distance correctly works when the inputs are the same array.
```
>>> np.isclose(distance(make_array(1, 2, 3), make_array(2, 4, 5)), 3)
True
```
failure_message: ❌ distance doesn't correctly work for our test input arrays.
success_message: ✅ distance correctly works for our test input arrays.


Tests for task_3_1_1 (points: 1)
```
>>> bottom_left in set([1, 2, 3, 4, 5])
True
```
failure_message: ❌ bottom_left should be 1, 2, 3, 4, or 5.
success_message: ✅ bottom_left is 1, 2, 3, 4, or 5.


Tests for task_3_1_2 (points: 1)
```
>>> bottom_right in set([1, 2, 3, 4, 5])
True
```
failure_message: ❌ bottom_right should be 1, 2, 3, 4, or 5.
success_message: ✅ bottom_right is 1, 2, 3, 4, or 5.


Tests for task_3_1_3 (points: 1)
```
>>> top_right in set([1, 2, 3, 4, 5])
True
```
failure_message: ❌ top_right should be 1, 2, 3, 4, or 5.
success_message: ✅ top_right is 1, 2, 3, 4, or 5.


Tests for task_3_1_4 (points: 1)
```
>>> top_left in set([1, 2, 3, 4, 5])
True
```
failure_message: ❌ top_left should be 1, 2, 3, 4, or 5.
success_message: ✅ top_left is 1, 2, 3, 4, or 5.


Tests for task_3_1_5 (points: 1)
```
>>> movie_genre_guess in set([1, 2])
True
```
failure_message: ❌ movie_genre_guess should be 1 or 2.
success_message: ✅ movie_genre_guess is 1 or 2.


Tests for task_3_1_6 (points: 5)
```
>>> len(my_features) >= 10
True
```
failure_message: ❌ my_features should contain at least 10 items.
success_message: ✅ my_features contains at least 10 items.
```
>>> np.all([f in test_movies.labels for f in my_features])
True
```
failure_message: ❌ my_features doesn't contain features from test_movies.
success_message: ✅ my_features contains features from test_movies.
```
>>> len(set(my_features)) >= 10
True
```
failure_message: ❌ my_features shouldn't contain duplicate words. 
success_message: ✅ my_features doesn't contain duplicate words.
```
>>> train_f = train_movies.select(my_features)
>>> np.count_nonzero(train_f.apply(lambda r: np.sum(np.abs(np.array(list(r)))) == 0)) < len(my_features)
True
```
failure_message: ❌ It looks like there are many movies in the training set that don't have any of your chosen words.  That will make your classifier perform very poorly in some cases.  Try choosing at least 1 common word.
success_message: ✅ It seems like you have at least one common word from the training set.
```
>>> test_f = test_movies.select(my_features)
>>> np.count_nonzero(test_f.apply(lambda r: np.sum(np.abs(np.array(list(r)))) == 0)) < 5
True
```
failure_message: ❌ It looks like there are many movies in the test set that don't have any of your chosen words.  That will make your classifier perform very poorly in some cases.  Try choosing at least 1 common word.
success_message: ✅ It seems like you have at least one common word from the test set.


Tests for task_3_1_7 (points: 4)
```
>>> genre_and_distances.labels == ('Genre', 'Distance')
True
```
failure_message: ❌ enre_and_distances doesn't contain the correct column labels.
success_message: ✅ genre_and_distances contains the correct column labels.
```
>>> genre_and_distances.num_rows == train_movies.num_rows
True
```
failure_message: ❌ genre_and_distances doesn't have the correct number of rows.
success_message: ✅ genre_and_distances has the correct number of rows.
```
>>> print(genre_and_distances.group('Genre'))
Genre    | count
comedy   | 102
thriller | 181
```
failure_message: ❌ genre_and_distances doesn't show the correct genres.
success_message: ✅ genre_and_distances shows the correct genres.
```
>>> np.allclose(genre_and_distances.column('Distance'), sorted(fast_distances(test_my_features.row(0), train_my_features)))
True
```
failure_message: ❌ genre_and_distances doesn't seem to have the correct distances.
success_message: ✅ genre_and_distances seems to have the correct distances.


Tests for task_3_1_8 (points: 2)
```
>>> genre_and_distances.take(np.arange(7)).group('Genre').index_by('Genre')[my_assigned_genre][0].item('count') >= 4
True
```
failure_message: ❌ my_assigned_genre doesn't to reflect the most common genre.
success_message: ✅ my_assigned_genre seems to reflect the most common genre.
```
>>> my_assigned_genre_was_correct == (my_assigned_genre == 'thriller')
False
```
failure_message: ❌ my_assigned_genre_was_correct doesn't reflect the correct value.
success_message: ✅ my_assigned_genre_was_correct reflects the correct value.


Tests for task_3_2_1 (points: 2)
```
>>> from collections import Counter
>>> g = train_movies.column('Genre')
>>> 
>>> def check(r, k):
...     t = test_my_features.row(r)
...     return classify(t, train_my_features, g, k) == Counter(np.take(g, np.argsort(fast_distances(t, train_my_features))[:k])).most_common(1)[0][0]
>>> check_5_nn = [check(i, 5) for i in np.arange(11)]
>>> all(check_5_nn)
True
```
failure_message: ❌ classify didn't work for our test case with k = 5.
success_message: ✅ classify worked for our test case with k = 5.
```
>>> from collections import Counter
>>> g = train_movies.column('Genre')
>>> 
>>> def check(r, k):
...     t = test_my_features.row(r)
...     return classify(t, train_my_features, g, k) == Counter(np.take(g, np.argsort(fast_distances(t, train_my_features))[:k])).most_common(1)[0][0]
>>> check_11_nn = [check(i, 11) for i in np.arange(11)]
>>> all(check_11_nn)
True
```
failure_message: ❌ classify didn't work for our test case with k = 11.
success_message: ✅ classify worked for our test case with k = 11.


Tests for task_3_2_2 (points: 1)
```
>>> from collections import Counter
>>> g = train_movies.column('Genre')
>>> r = np.where(test_movies['Title'] == 'godzilla')[0][0]
>>> t = test_my_features.row(r)
>>> godzilla_expected_genre = Counter(np.take(g, np.argsort(fast_distances(t, train_my_features))[:15])).most_common(1)[0][0]
>>> godzilla_genre == godzilla_expected_genre
True
```
failure_message: ❌ godzilla_genre doesn't seem to reflect the correct response.
success_message: ✅ godzilla_genre seems to reflect the correct response.


Tests for task_3_2_3 (points: 1)
```
>>> def check(r):
...     t = test_my_features.row(r)
...     return classify(t, train_my_features, train_movies.column('Genre'), 15) == classify_feature_row(t)
>>> all([check(i) for i in np.arange(15)])
True
```
failure_message: ❌ classify_feature_row doesn't seem to work correctly with our test case.
success_message: ✅ classify_feature_row worked correctly with our test case.


Tests for task_3_3_1 (points: 2)
```
>>> 0 <= accuracy <= 1
True
```
failure_message: ❌ accuracy should be a value between 0 and 1, inclusive.
success_message: ✅ accuracy is a value between 0 and 1, inclusive.


Tests for task_3_3_2 (points: 3)
```
>>> test_movie_correctness.labels == ('Title', 'Genre', 'Was correct')
True
```
failure_message: ❌ test_movie_correctness doesn't have the correct column labels.
success_message: ✅ test_movie_correctness has the correct column labels.
```
>>> test_movie_correctness.num_rows == test_movies.num_rows
True
```
failure_message: ❌ test_movie_correctness doesn't have the correct number of rows.
success_message: ✅ test_movie_correctness has the correct number of rows.
```
>>> print(test_movie_correctness.group('Genre'))
Genre    | count
comedy   | 17
thriller | 33
```
failure_message: ❌ test_movie_correctness shouldn't modify the original test_movies table.
success_message: ✅ test_movie_correctness doesn't seem to modify the original test_movies table.
